Economic and Social Modulations of Innate Decision-Making in Mice Exposed to Visual Threats

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

This study by Li and colleagues examines how defensive responses to visual threats during foraging are modulated by both reward level and social hierarchy. Using a naturalistic paradigm, the authors test how the availability of water or sucrose, with sucrose being more rewarding than water, shapes escape behavior in mice exposed to looming stimuli of different intensities, which are used to probe perceived threat level and defensive responses. In parallel, the study compares dominant and subordinate animals to assess how social rank biases the trade off between reward seeking and threat avoidance. By combining detailed behavioral analyses with computational modeling, the work addresses how reward level and social context jointly influence escape decisions in an ethologically relevant setting.

Across the different experimental conditions, perceived threat level is the main determinant of behavior. The authors show that looming stimuli associated with higher threat (contrast) consistently elicit faster and more robust escape responses than lower threat stimuli. This effect is particularly evident during early exposures, when animals are highly vigilant and have not yet habituated to the looming stimulus (learned that it is not dangerous). Later they described that as animals gain experience and habituate, behavior becomes more flexible, and reward level begins to exert a graded modulation of the escape response. Importantly, the authors show that under high threat conditions increasing reward value leads to more frequent and faster escape rather than greater reward pursuit. This finding is particularly relevant, as it suggests that highly valued rewards can heighten vigilance and thereby enhance responsiveness to threat, highlighting that reward does not simply compete with defensive behavior but can also reshape it depending on the perceived level of danger, in contrast to low threat conditions, where threat can be more easily outweighed by reward. Thus, an important conceptual contribution of the study is the introduction of vigilance as a useful framework to interpret these effects. Vigilance is treated as a behavioral state reflecting heightened attention to potential danger. In line with what is known from natural foraging, mice initially maintain high vigilance when confronted with an innate threat. This perspective helps clarify a finding that might otherwise appear counterintuitive. One might expect higher rewards to motivate animals to tolerate risk, explore more, and habituate faster in any scenario. Instead, the data suggest that highly rewarding outcomes can elevate vigilance, making animals more responsive to threat and leading to faster or more frequent escape under high threat conditions. In this sense, reward does not simply compete with threat but can also amplify sensitivity to it, depending on the internal state of the animal.

The social results are particularly interesting in this context as well. Dominant mice consistently prioritize avoidance over reward, showing stronger escape responses and slower habituation than subordinates. This behavior is well captured by the vigilance framework proposed by the authors: dominant animals appear to maintain higher vigilance, which biases decisions toward threat avoidance. The authors further suggest that stable social relationships sustain high vigilance and slow habituation, framing this as an evolutionarily conserved strategy that may enhance survival. This interpretation provides a valuable perspective on how social structure shapes defensive behavior beyond immediate physical interactions. At the same time, there are important limitations to this interpretation. All experiments were conducted in male mice, and it is possible that the relationship between social hierarchy, vigilance, and defensive behavior would differ substantially in females. In addition, the idea that stable social relationships maintain elevated vigilance does not straightforwardly align with broader views of social stability as protective for mental health and as a buffer against anxiety and stress. These points do not undermine the findings but suggest that the social effects described here should be interpreted with caution and within the specific context of the task and sex studied.

We thank the reviewer for raising this important point. In the context of repeated looming exposure, slower habituation reflects more sustained vigilance over time. Compared to individually housed mice, group-housed mice exhibit slower habituation (Lenz et al., 2022), and pair-housed mice showed even slower habituation in our current work. Importantly, this pattern does not indicate that pair-housed mice have higher overall vigilance than individually housed animals. Although individually housed mice habituate more quickly, they display higher initial vigilance, as reflected by their increased probability of escaping in response to looming stimuli (Lenz et al., 2022). Thus, pairhoused mice exhibited reduced defensive responses compared to individually housed animals, consistent with a social buffering effect.

Furthermore, in a separate study (Rank- and Threat-Dependent Social Modulation of Innate Defensive Behaviors; Li, Gao, Li, 2026, eLife 15:RP109571), we directly compared responses to looming stimuli when mice were tested alone versus in the presence of a social partner and observed clear evidence of social buffering.

Another important limitation is that the neural mechanisms underlying these effects remain speculative. The manuscript includes an extensive discussion of candidate circuits, particularly involving the superior colliculus and downstream structures, but this section is necessarily based on prior literature rather than on data presented in the study. Given the complexity of the circuits involved in integrating internal state, reward, social context, and vigilance, the current work should be viewed as providing a strong behavioral and conceptual framework rather than direct insight into underlying neural mechanisms.

We fully agree that the proposed neural mechanisms remain speculative and that the circuits involved in integrating internal state, reward, and social context are likely far more complex. We have revised the manuscript to acknowledge this limitation.

Methodologically, the behavioral paradigm is well suited for studying escape decisions in socially housed animals, and the machine learning based classification of defensive responses is a clear strength. The computational model provides a useful formalization of how threat level, reward level, and vigilance interact and may be valuable for other laboratories studying escape, approach avoidance, or conflict situations, particularly as a way to classify behavioral outcomes after pose estimation. More generally, the work will be of interest to the neuroethology community for its detailed characterization of escape behavior under naturalistic conditions.

Given the ethological nature of the study and the high inter individual variability reported by the authors, clarity and precision in the methods are especially important for reproducibility. While the revised manuscript addresses many earlier concerns, some aspects remain slightly difficult to follow. For example, the main text states that animals were not water deprived to avoid differences in internal state, whereas parts of the methods describe conditions in which animals were water deprived, suggesting that internal state manipulation may differ across experiments. Clearer separation and explanation of these conditions would further strengthen confidence in the work.

To improve clarity, we have revised the Methods section to clearly distinguish between experimental conditions that involved water deprivation and those that did not.

Overall, this study provides a rich and thoughtful analysis of how reward level and social hierarchy modulate defensive behavior through changes in vigilance. It offers a useful conceptual advance for thinking about escape behavior in naturalistic settings and lays a solid foundation for future work aimed at linking these behavioral states to underlying neural circuits.

Reviewer #2 (Public review):

Zhe Li and colleagues investigate how mice exposed to visual threats and rewards balance their decisions in favour of consuming rewards or engaging in defensive actions. By varying threat intensity and reward value, they first confirm previous findings showing that defensive responses increase with threat intensity and that there is habituation to the threat stimulus. They then find that water-deprived mice have a reduced probability of escaping from low contrast visual looming stimuli when water or sucrose are offered in the environment, but that when the stimulus contrast is high, the presence of sucrose or water increases the probability of escape. By analysing behaviour metrics such as the latency to flee from the threat stimulus, they suggest that this increase in threat sensitivity is due to increased vigilance. Analysis of this behaviour as a function of social hierarchy shows that dominant mice have higher threat sensitivity, which is also interpreted as being due to increased vigilance. These results are captured by a drift diffusion model variant that incorporates threat intensity and reward value.

The main contribution of this work is quantifying how the presence of water or sucrose in water-deprived mice affects escape behaviour. The differential effects of reward between the low and high contrast conditions are intriguing, but I find the interpretation that vigilance plays a major in this process not supported by the data. The idea that reward value exerts some form of graded modulation of the escape response is also not supported by the data. In addition, there is very limited methodological information, which makes assessing the quality of some of the analyses difficult, and there is no quantification on the quality of the model fits.

(1) The main measure of vigilance in this work is reaction time. While reaction time can indeed be affected by vigilance, reaction times can vary as a function of many variables, and be different for the same level of vigilance. For example, a primate performing the random dot motion task exhibits differences in reaction times that can be explained entirely by the stimulus strength. Reaction time is therefore not a sound measure of vigilance, and if a goal of this work is to investigate this parameter, then it should be measured. There is some attempt at doing this for a subset of the data in Figure 3H, by looking at differences in the action of monitoring the visual field (presumably a rearing motion, though this is not described) between the first and second trials in the presence of sucrose. I find this an extremely contrived measure. What is the rationale for analysing only the difference between the first and second trials? Also, the results are only statistically significant because the first trial in the sucrose condition happens to have zero up action bouts, in contrast to all other conditions. I am afraid that the statistics are not solid here. When analysing the effects of dominance, a vigilance metric is the time spent in the reward zone. Why is this a measure of vigilance? More generally, measuring vigilance of threats in mice requires monitoring the position of the eyes, which previous work has shown is biased to the upper visual field, consistent with the threat ecology of rodents.

(2) In both low and high contrast conditions, there are differences in escape behaviour between no reward and water or sucrose presence, but no statistically significant differences between water and sucrose (eg: Figure 3B). I therefore find that statements about reward value are not supported by the data, which only show differences between the presence or absence of reward. Furthermore, there is a confound in these experiments, because according to the methods, mice in the no-reward condition were not water-deprived. It is thus possible that the differences in behaviour arise from differences in the underlying state.

(3) There is very little methodological information on behavioural quantification. For example, what is hiding latency?

Is this the same are reaction time? Time to reach the safe zone? What exactly is distance fled? I don't understand how this can vary between 20 and 100cm. Presumably, the 20cm flights don't reach the safe place, since the threat is roughly at the same location for each trial? How is the end of a flight determined? How is duration measured in reward zone measures, e.g., from when to when? How is fleeing onset determined?

(4) There is little methodological information on how the model was fit (for example, it is surprising that in the no reward condition, the r parameter is exactly 0. What this constrained in any way), and none of the fit parameters have uncertainty measures so it is not possible to assess whether there are actually any differences in parameters that are statistically significant.

These are the public reviews for the original submission. The corresponding authors responses are provided below.

(1) We agree that reaction time can be influenced by multiple factors, including stimulus strength. Consistent with this, reaction times (i.e. latencies to flee) were substantially shorter under high-contrast conditions (Figure 3E). However, even under the same high-contrast condition, reaction times were significantly shorter in the water condition compared to the no-reward condition, suggesting that other factors such as vigilance may contribute.

Upward-directed attention includes rearing, up-stretching, and upward head orientation, which will be clarified in the Method section. To address concerns about statistical validity, we will quantify these behaviors across the first 10 trials rather than limiting the analysis to the first two.

As for the dominance-related results, we interpret them as reflecting both enhanced vigilance and reduced reward-seeking behavior. Time spent in the reward zone is not a measure of vigilance but an indicator of reward-seeking motivation. We will clarify this in the revised manuscript.

(2) In Figure 3B, the difference between water and sucrose conditions did not reach statistical significance (p = 0.08). We plan to collect additional data to determine whether this is due to limited statistical power. It is also possible that some behavioral readouts are more sensitive to the differences between water and sucrose conditions. For example, Figure 3F shows that escape speed was significantly higher in the sucrose than in the water condition under high-contrast stimulation.

Thank you for pointing this out. To control for the potential confounds related to internal state, mice were not water-deprived under any of the three conditions in Figures 3A-3H. We will clarify this in the main text and Methods. For Figures 3I-3M, which compare decision-making under no-reward and water conditions, we will conduct additional experiments using non-deprived mice in the water condition.

(3) Hiding latency was defined as the time from stimulus onset to the animal’s arrival at the safe zone. Reaction time was quantified as the latency to flee, measured from stimulus onset to the initiation of the first flight state. The flight state was defined as locomotion exceeding 10 cm at a speed greater than 10 cm/s. Distance fled was defined as the distance covered between stimulus onset and offset for all trials. However, in trials classified as no reaction or freezing, this measure does not accurately reflect escape behavior. We will therefore rename it as distance under threat to better capture its meaning. The reward zone was defined as the region within 15 cm of the reward port at the end of the arena. Duration in the reward zone was measured as the time spent within this region during the 20 seconds following stimulus onset. In Figure 4E, the percentage of time spent in the reward zone was calculated relative to the total time the mouse remained in the arena during the 2-hour social session.

All definitions and additional details on behavioral quantification will be included in the revised Methods section.

(4) We appreciate the comment and agree that further clarification is needed. We will provide a more detailed description of the model fitting procedure in the revised Methods section. Specifically, the drift rate parameter (r), which reflects the perceived reward value, was constrained to zero in the no-reward condition. To enable statistical comparison across conditions, we will report uncertainty measures for all fit parameters.

Comments on the revised manuscript:

The manuscript has been revised and improved significantly by the addition of methodological details and new analysis. I remain, however, unconvinced by the argument that increased vigilance in the presence of reward leads to heightened escape behaviour.

In response to my criticism that the work does not measure vigilance directly, the authors have included measures of foraging interval and foraging speed, which they state are "two direct behavioral analyses of vigilance". I disagree - like reaction time, foraging speed and foraging interval can be modulated, for example, by changes in threat sensitivity. Increased threat sensitivity comes with diverse behavioral changes that may well include increased vigilance, but foraging interval and foraging speed can certainly change without the animal expressing increased vigilance behaviors. A bigger issue I still have though, is with the conclusion that the presence of reward increases "direct escape behaviors". Comparing the no reward, water and sucrose groups indeed shows a difference (which is now clear after the split into early and late phases), but the issue is that these are different mice. As the text is written, is sounds like introducing reward will acutely increase escape. But if we look at the raw data show in Figure 2C, what I think is happening is that the presence of reward is decreasing habituation to the stimulus. The data for trials 1 and 10 in the three conditions show this - there is habituation with no reward (reaction times are all shifting to the right), a bit less with water and very little with sucrose. This is interesting in its own right and we can speculate why it might be happening, but I think this is conceptually different from what the authors are proposing.

We agree that vigilance is not directly observable as a single variable. Our intent was not to claim that foraging speed and foraging interval provide a direct measure of vigilance, but rather to suggest that they may serve as indirect behavioral correlates.

We also considered an alternative interpretation: these two measures could reflect perceived reward value under high-threat conditions across distinct reward types. If that were the case, animals would be expected to exhibit shorter intervals and faster speeds across no reward, water, and sucrose conditions. However, our data do not support this interpretation (Figures 3L and 3M), suggesting that these measures are more likely correlated with vigilance.

Furthermore, it is unlikely that changes in foraging interval and speed are driven by altered threat sensitivity, as animals could not see the threat during most of the foraging bout and only encountered it at the end.

Regarding the conclusion that the presence of reward increases direct escape behaviors, our interpretation is that increased reward value reduces habituation, thereby maintaining higher vigilance during the late phase. This was discussed in the second-to-last paragraph of the "Economic and social modulations of innate decision-making under threat" subsection in the Discussion.

Reviewer #3 (Public review):

Male mice were tested in a classic behavioral "flee the looming stimulus" paradigm. This is a purely behavioral study; no neural analyses were done. Mice were housed socially, but faced the looming stimulus individually, using an elegant automated tunnel (see videos for clarity).

The additional changes made to the paper clarify the work done. While there are some limitations (male mice, weird stimulus), the general results are interesting and a valuable addition to the experimental literature. The main claim of the paper is that the different rewards (none, water, sucrose) did not change the escape properties early in learning, but did late, particularly that in the late (already experienced) conditions, reward value (assuming sucrose > water > no reward) interacted with the salience of the looming stimulus (light gray, dark gray). (Panels 3D, 3G, 3K, 3N).

For readers, I want to note that one of the most interesting results is actually in Figure S2, where they find that a looming stimulus behind the mouse still makes a mouse run to the nest. In these conditions, the mouse runs past the looming stimulus to get to safety! (I also do love the video of the mouse running around the barriers like a snake to get home.)

I have a few minor clarification questions and a few notes that I think would be useful additions for authors and readers to think about.

Dominance: What does the mouse social science literature say about the "test tube" test? What can we conclude from this test? This would be useful when trying to understand what is causing the dominance/submissive difference in responses. Figure 4 shows that the dominant mice are more risk-averse than the submissive mice. Is "dominance" in the test-tube actually a measure of risk-seeking? Is the issue that the submissive mice don't think they can get back to the food-site easily, so they are less willing to sacrifice the current (if dangerous) foraging opportunity? Is the issue that the submissive mice can't get back to the nest? As I understand it, the nest was always available to all the mice, so I suspect inability to get to the nest is an unlikely hypotheses. Is the issue that the submissive mice also don't feel safe in the nest?

The tube test is a widely used assay in the rodent social behavior literature to assess dominance hierarchies, operationally defined by the ability of one animal to force its opponent to retreat from a narrow tube. Importantly, this assay does not directly measure risk-seeking or anxiety-related traits, but rather competitive outcomes during social conflict. Furthermore, our data indicate that the behavioral responses of subordinate mice to looming stimuli are primarily driven by the visual threat itself rather than by social avoidance. This point was elaborated in the second paragraph of the “Social modulation of innate decision-making” subsection in the Results section.

Limitations of the study: There is an acknowledged limitation to male mice, and the limitations of the small data sets that are typical of such experiments. In addition, however, it is also worth noting the strangeness of the looming stimulus, which is revealed clearly in the videos. The stimulus is a repeating growing circle, growing in a single location within the environment. The stimulus repeats 10 times, once per second. This is not what an attacking hawk or owl would look like. (I now have this image of an owl diving down, and then teleporting up and diving down again.) Note - I am fine with this stimulus. It produces an interesting experiment and interesting results. I do not think the authors need to change anything in their paper, but readers need to recognize that this is not a "looming predator".

These "limitations" are better seen as "caveats" when folding these results in with the rest of the literature that has gone before and the literature to come. (Generally, I do not believe that science works by studies making discoveries that change how we think about problems - instead, science works by studies adding to the literature that we integrate in with the rest of the literature.) Thus, these caveats should not be taken as problems with the study or as fixes that need to be done. Instead, they are notes for future researchers to notice if differences are found in any future studies.

Thus, my only suggestion is that I think authors could write a more careful paper by using the past and subjunctive tense appropriately. Experimental observations should be in past tense, as in "the influence of reward was contextdependent and emerged in the late phase" instead of "the influence of reward is context-dependent and emerges in the late phase" - it emerged in the late phase this once - it might not in future experiments, not due to any fault in this experiment nor due to replicability problems, but rather due to unexpected differences between this and those future experiments. At which point, it will be up to those future experiments to determine the difference. Similarly, large conclusions should be in the subjunctive tense, as in "these data suggest that threat intensity is likely to be the primary determinant of decision making" rather than "threat intensity is the primary determinant of decision making", because those are hypotheses not facts.

We thank the reviewer for the helpful suggestions and have revised the Abstract accordingly.

Recommendations for the authors:

Reviewer #3 (Recommendations for the authors):

Figure 5: The points in panel 5G and 5H are unreadable. What are these stars and symbols supposed to mean? They are also too small to see without zooming way in.

We have increased the symbol size.

Figure 5: What is the final panel of 5J? I did not understand this panel at all. The first three panels of 5J (threat-based detection, reward-based detection, vigilance-based detection) are, I believe, three patterns we should look for in the data. But then what is the "experimental results" section? It contains all three, but they don't overlap? Shouldn't we have an experimental results section for each condition?

Panel 5J was to compare three hypothesized decision patterns with the experimentally observed data. To make this distinction explicit, we have revised the panel titles to: “H1: Threat-based decisions,” “H2: Reward-based decisions,” “H3: Vigilance-based decisions,” and “Experimental results.”

Thank you for including the videos. They made the task construction and the stimulus much clearer.

Changes from the earlier version to the revised version

Abstract

Earlier version: When confronted by predators, animals make innate decisions with rapid reaction times—a trait shaped by natural selection to maximize survival. However, rapid reactions are effective only when grounded in accurate judgments and appropriate choices, which often require cognitive control. To address how such choices are shaped, we developed a behavioral paradigm to investigate how threat intensity, reward value, and social hierarchy influence decision-making in foraging mice exposed to overhead visual threats. Using a machine learning-based approach, we classified defensive responses into four distinct decision types. We observed rapid habituation to repeated looming threats, with substantial inter-individual variability in the rate of habituation. Across both early and late phases of habituation, threat intensity is the primary determinant of decision-making, strongly driving behavior towards escape. In contrast, the influence of reward is context-dependent and emerges primarily in the late phase: under low-threat conditions, higher rewards suppress defensive responses, consistent with value-based decision theory, whereas under high-threat conditions, higher rewards promote escape, potentially reflecting heightened vigilance. Innate decision-making is further modulated by social hierarchy, with dominant mice showing greater vigilance and a stronger bias towards risk-averse behaviors, while subordinates are more reward-driven. To understand the underlying decision-making process, we developed a drift-diffusion leaky integrator model that successfully captures how threat intensity, reward value, and vigilance are integrated. Together, these findings reveal the economic and social modulation of innate decisions, offering insights into the computational mechanisms underlying the interplay between instinctive reactions and cognitive control.

Revised version: When confronted by predators, animals make innate decisions with rapid reaction times—a trait shaped by natural selection to maximize survival. However, rapid reactions are effective only when grounded in accurate judgments and appropriate choices, which often require cognitive control. To address how such choices are shaped, we developed a behavioral paradigm to investigate how threat intensity, reward value, and social hierarchy shape decision-making in foraging mice exposed to overhead visual threats. Using a machine learning-based approach, we classified defensive responses into four distinct decision types. Mice showed rapid habitua-tion to repeated looming threats, with substantial inter-individual variability in the rate of habituation. Across both early and late phases of habituation, threat intensity emerged as the primary determinant of decision-making, strongly driving behavior toward escape. In contrast, the influence of reward was context-dependent and became evident primarily in the late phase: under low-threat conditions, higher re-wards suppressed defensive responses, consistent with value-based decision theory; whereas under high-threat conditions, higher rewards promoted escape, potentially reflecting heightened vigilance. Innate decision-making was further modulated by social hierarchy, with dominant mice showing greater vigilance and a stronger bias toward risk-averse behaviors, while subordinates were more reward-driven. To understand the underlying decision-making process, we developed a drift-diffusion leaky integrator model that successfully captures how threat intensity, reward value, and vigilance are integrated. Together, these findings reveal how economic and social factors modulate innate decisions and provide a computational framework for understanding the interplay between instinctive reactions and cognitive control.

Introduction

How animals make decisions among alternative actions is a key question in neuroscience. In natural environments, decisions arise from the integration of sensory inputs, internal states, and learned experience (Meister, 2022). In the laboratory, both learned and innate paradigms have been used to study decision-making. Some learned paradigms, such as two-alternative forced choice (2-AFC) and GO-NOGO tasks (Andermann et al., 2010; Burgess et al., 2017), often require weeks of training in head-fixed animals; whereas other learned paradigms require less training over several days in freely moving animals, including maze exploration (Barnes, 1979; Morris, 1984; Rosenberg et al., 2021; Small, 1901), foraging decisions (Hayden, 2018; Hayden et al., 2011; Steiner and Redish, 2014), and active avoidance (Bravo-Rivera et al., 2014; Moscarello and LeDoux, 2013; Mowrer and Lamoreaux, 1946). While these paradigms have significantly advanced our understanding of the neural mechanisms underlying learned decisions, the training itself may engage neural circuits that differ from those evolved for innate decision-making (Steinmetz et al., 2019).

Earlier version: In contrast, innate decisions—such as whether to escape from an approaching aerial predator (De Franceschi et al., 2016; Yilmaz and Meister, 2013) or a simulated ground threat like a robogator (Amir et al., 2015; Choi and Kim, 2010)—are executed without prior training and may more directly reflect the function of neural circuits shaped by natural selection. In complex environments, making correct and rapid innate defensive decisions often requires cognitive control to assess risk and evaluate alternative defensive strategies (Evans et al., 2019). Defensive responses to terrestrial predators have been well studied: escape decisions depend on predator–prey distance (Ydenberg and Dill, 1986), and the predatory imminence continuum theory describes a graded set of defensive behaviors tuned to perceived threat level (Fanselow and Lester, 1988). Responses to approaching aerial predators are likely also scaled with the perceived threat intensity, which depends on the physical properties of the threat (De Franceschi et al., 2016; Liden and Herberholz, 2008; Tammero and Dickinson, 2002; Yang et al., 2020; Yilmaz and Meister, 2013), prior experience (Vale et al., 2017), and environmental context, including housing conditions (Lenzi et al., 2022). At the neural level, the superior colliculus (SC) is a central node for processing looming-evoked defensive responses, with distinct output pathways mediating escape versus freezing (Evans et al., 2018; Shang et al., 2018; Wei et al., 2015; Zhou et al., 2019).

Revised version: In contrast, innate decisions—such as whether to escape from an approaching aerial preda-tor (De Franceschi et al., 2016; Yilmaz and Meister, 2013) or a simulated ground threat like a robogator (Amir et al., 2015; Choi and Kim, 2010)—are executed without prior training and may more directly reflect the function of neural circuits shaped by natural selection. In complex envi-ronments, making correct and rapid innate defensive decisions often requires cognitive control to assess risk and evaluate alternative defensive strategies (Evans et al., 2019). Defensive responses to terrestrial predators depend on predator–prey distance (Ydenberg and Dill, 1986) and are well described by the predatory imminence continuum theory (Fanselow and Lester, 1988). Responses to approaching aerial predators are likely scaled similarly according to the perceived threat intensity, which is influenced by the physical properties of the threat (De Franceschi et al., 2016; Liden and Herberholz, 2008; Tammero and Dickinson, 2002; Yang et al., 2020; Yilmaz and Meister, 2013), prior experience (Vale et al., 2017), and environmental context, including housing conditions (Lenzi et al., 2022). At the neural level, the superior colliculus (SC) serves as a central hub for processing looming-evoked defensive responses, with distinct output pathways mediating escape versus freezing behaviors (Evans et al., 2018; Shang et al., 2018; Wei et al., 2015; Zhou et al., 2019).

Earlier version: Despite this progress, important gaps remain. Mice are social animals, and most predator encounters occur during foraging, yet it is unclear how they weigh perceived risks and rewards when making defensive decisions, or how these decisions are influenced by social hierarchy. To address these questions, we developed an ecologically relevant behavioral paradigm to investigate decision-making in foraging mice exposed to overhead visual threats. Our findings demonstrate that defensive choices are jointly influenced by threat intensity, reward magnitude, and social hierarchy, providing a behavioral framework for future investigations into the neural mechanisms underlying cognitive control of innate decision-making.

Revised version: Despite this progress, important gaps remain. Mice are social animals, and most predator encounters occur during foraging, yet it is unclear how they weigh perceived threat and reward when making defensive decisions, or how these decisions are influenced by social hierarchy. To address these questions, we developed an ecologically relevant behavioral paradigm to investigate decision-making in foraging mice exposed to overhead visual threats. Our findings demonstrate that defensive choices are jointly influenced by threat intensity, reward magnitude, and social hierarchy, providing a behavioral framework for future investigations into the neural mechanisms underlying cognitive control of innate decision-making.

Results

A behavioral paradigm for studying innate decision-making in mice

Earlier version: To investigate how animals make innate decisions in natural environments, we designed a behavioral paradigm to simulate the defensive behavior of foraging mice in the wild. In this paradigm, a group of 2–5 co-housed mice was placed in a nest, with each individual identified using a radio frequency identification (RFID) tag. Only one mouse was allowed to enter a linear arena to receive the reward delivered at the end. As the mouse approached the reward, an overhead expanding dark disc that mimics the approach of an aerial predator was triggered (Figure 1A), forcing the animal to decide whether to risk getting the reward or defend itself for safety. Behavioral data were recorded using a ground-mounted camera, and DeepLabCut was used to track the movements of the mouse’s nose and tail base (Mathis et al., 2018) (Figure S1A). Compared to the exploratory phase, the presence of looming stimuli significantly increased the frequency of arena entries but decreased the duration of each visit (Figures 1B and 1C).

Revised version: To investigate how animals make innate decisions in natural environments, we designed a behavioral paradigm to simulate the defensive behavior of foraging mice in the wild. In this paradigm, a group of 2–5 co-housed mice was placed in a nest, with each individual identified using a radio frequency identification (RFID) tag. Only one mouse was allowed to enter a linear arena to receive the reward delivered at the end. As the mouse approached the reward, an overhead expanding dark disc that mimics the approach of an aerial predator was triggered (Figure 1A), forcing the animal to decide whether to risk getting the reward or to seek safety in the nest. Behavioral data were recorded using a ground-mounted camera, and DeepLabCut was used to track the movements of the mouse’s nose and tail base (Mathis et al., 2018) (Figure S1A). Compared to the exploratory phase, the presence of looming stimuli significantly increased the frequency of arena entries but decreased the duration of each visit (Figures 1B–C).

A behavioral paradigm for investigating innate decision-making in mice.

Earlier version: (A) Schematic of the behavioral assay (3D and top-down views). (B) Top: arena occupancy patterns for five example mice in one session. Bottom: visit frequency and average duration for each visit. Colors mark the mouse’s identity. (C) Visit frequency and duration under exploration and threat conditions. Error bars represent the standard deviation. N = 46 mice, paired t-test. (D) The pipeline for behavioral classification. (E) Distribution of decisions across 3862 trials from 140 mice. (F) Left: distance to the safe zone over time for four decision types (10 example trials each). Red dashed lines mark the onset of each stimulus repetition; solid lines mark the end of the last repetition. Grey shade indicates the safe zone. Right: locomotion speed towards the safe zone over time for the same trials. Positive speed indicates movement towards the safe zone. (G) Distribution of the longest stationary time for “freezing” and “assessment+escape”. N = 224 and 1667 trials, Mann-Whitney-Wilcoxon test. Blue lines mark the median; red solid lines mark the end of the last repetition. (H) Distribution of latency to flee for “escape” and “assessment+escape”. N = 458 and 1667 trials, Mann-Whitney-Wilcoxon test. Blue lines mark the median; red dashed and solid lines mark the end of the first and last repetitions, respectively. (I) Distribution of peak speed for “escape” and “assessment+escape”. N = 458 and 1667, Mann-Whitney-Wilcoxon test. (J) Distribution of hiding latency in the safe zone for “escape” and “assessment+escape”. N = 455 and 1563, Mann-Whitney-Wilcoxon test. (K) Distribution of post-decision behavioral states. (L) Distribution of fear recovery time for “escape”, “assessment+escape”, and “freezing”. N = 74, 458, and 205, Kruskal-Wallis test followed by Dunn’s post-hoc test with Holm correction. *p < 0.05, **p < 0.01, ***p < 0.001.

Revised version: (A) Schematic of the behavioral assay (3D and top-down views). (B) Top: arena occupancy patterns for five example mice in one session. Bottom: visit frequency and average duration for each visit. Different colors mark the mouse’s identity. (C) Visit frequency and duration under exploration and threat conditions. Error bars represent standard deviation. n = 46 mice, paired t-test. (D) Pipeline for behavioral classification. (E) Distribution of decisions across 3862 trials from 140 mice. (F) Left: distance to the safe zone over time for four decision types (10 example trials each). Red dashed lines mark the onset of each stimulus repetition; solid lines mark the end of the last repetition. Gray shading indicates the safe zone. Right: locomotion speed toward the safe zone over time for the same trials. Positive speed indicates movement toward the safe zone. (G) Distribution of the longest stationary time for “freezing” and “assessment+escape”. n = 224 and 1667 trials, Mann–Whitney–Wilcoxon test. Blue lines mark the median; red solid lines mark the end of the last repetition. (H) Distribution of latency to flee for “direct escape” and “assessment+escape”. n = 458 and 1667 trials, Mann–Whitney–Wilcoxon test. Blue lines mark the median; red dashed and solid lines mark the end of the first and last repetitions, respectively. (I) Distribution of peak speed for “direct escape” and “assessment+escape”. n = 458 and 1667, Mann–Whitney–Wilcoxon test. (J) Distribution of hiding latency in the safe zone for “direct escape” and “assessment+escape”. n = 455 and 1563, Mann–Whitney–Wilcoxon test. (K) Distribution of post-decision behavioral states. SZ, safe zone. (L) Distribution of fear recovery time for “direct escape”, “assessment+escape”, and “freezing”. n = 74, 458, and 205, Kruskal-Wallis test followed by Dunn’s post-hoc test with Holm correction. *p < 0.05, **p < 0.01, ***p < 0.001.

Earlier version: To identify distinct behavioral patterns in response to the looming stimuli, we defined 19 behavioral features from key body points of the mouse and fed them into a random forest classifier to predict its decisions (Figures 1D and S1B, see Methods and Materials). Animal decisions across 3861 trials were categorized into four types: direct escape (11.8%), escape after assessment (43.2%), freezing (5.8%), and no response (39.2%, Figures 1E and 1F). The classification model achieved an accuracy of 95% (Figure S1C). A key distinction between “assessment+escape” and “freezing” decisions was the duration of stationary behavior: mice that decided to freeze remained stationary significantly longer than those in the “assessment+escape” group (Figure 1G). Furthermore, the latency to flee differed significantly between “escape” and “assessment+escape” decisions (Figure 1H).

Revised version: To identify distinct behavioral patterns in response to looming stimuli, we defined 19 behavioral features from key body points of the mouse and fed them into a random forest classifier that predicted decisions with 95% accuracy (Figures 1D and S1B–C, see Materials and methods). Across 3862 trials, decisions were categorized into four types: direct escape (11.8%, Video 1), escape after assessment (assessment+escape, 43.2%, Video 2), freezing (5.8%, Video 3), and no response (39.2%, Video 4) (Figures 1E–F). A key distinction between assessment+escape and freezing decisions was the duration of stationary behavior: mice that decided to freeze remained stationary significantly longer than those in the assessment+escape group (Figure 1G). Furthermore, the latency to flee differed significantly between direct escape and assessment+escape decisions (Figure 1H).

Earlier version: We hypothesized that fear level increases progressively across the decision types: “freezing”, “escape after assessment”, and “escape”. This is supported by the observations that mice in the “escape” group exhibited higher flee speeds (Figure 1I) and shorter hiding latency (Figure 1J). Additionally, the proportion of mice that recovered within 20 seconds of the stimulus was about 20%, 40%, and 80% for “escape”, “assessment+escape”, and “freezing”, respectively (Figure 1K). Consistently, recovery time decreased progressively across these decision types (Figure 1L). Our findings align with the theory of predatory imminence continuum (Fanselow and Lester, 1988), which proposes that defensive responses become more intense as the perceived threat level increases.

Revised version: We hypothesized that fear level increases progressively across the decision types: freezing, assess-ment+escape, and direct escape. This hypothesis is supported by the observations that mice in the direct escape group exhibited higher escape speeds (Figure 1I) and shorter hiding latencies (Figure 1J). Additionally, the proportion of mice that recovered within 20 s after stimulus onset was about 20%, 40%, and 80% for direct escape, assessment+escape, and freezing, respectively (Figure 1K). Consistently, recovery time decreased progressively across these decision types (Figure 1L). Our findings align with the predatory imminence continuum theory (Fanselow and Lester, 1988), which proposes that defensive responses become more intense as the perceived threat level increases.

Mice make economic decisions modulated by vigilance

In the wild, most prey–predator encounters occur while prey animals are foraging, and prey have evolved to maintain high vigilance during foraging to enhance survival. Yet how food value, threat intensity, and vigilance interact to shape defensive decisions remains unclear. To address this question, we developed a behavioral assay that simulates how a foraging animal responds to an approaching aerial predator under different threat and reward conditions (Figure 2A). To maintain a consistent internal state across conditions, mice were not water-deprived. The higher reward value of sucrose over water was validated by measuring consumption during exploration (Figure S3A).

Mice learn quickly from experience.

Earlier version: (A) Schematic of the behavioral assay for studying the economic modulation of innate decision-making. (B-C) Distance to the safe zone and locomotion speed over time for the 1st and 10th trials in response to low- and high-contrast looming stimuli across different reward conditions. Right inset in (B) is the distribution of distance to the safe zone at the end of each trial. Dashed lines mark the start of each stimulus, and solid lines mark the stimulus offset. N = 11 (no reward, low), 13 (water, low), 10 (sucrose, low), 10 (no reward, high), 8 (water, high), 10 (sucrose, high) mice. (D) Summary of the decisions across the first 10 trials under different risk and reward conditions. N = 11, 13, 10, 10, 8, 10 mice. (E-H) Escape distance under threat, duration in reward zone, peak escape speed, and latency to flee across trials in all conditions. Shade denotes the standard error of the mean. (I) Transition trials marking the shift between phases across conditions. Red lines indicate the median; boxes span the interquartile range (IQR); whiskers extend to 1.5 × IQR beyond the box.

Revised version: (A) Schematic of the behavioral assay for studying the economic modulation of innate decision-making. (B-C) Distance to the safe zone and locomotion speed over time for the 1st and 10th trials in response to low- and high-contrast looming stimuli across different reward conditions. Right inset in (B) shows the distribution of distance to the safe zone at the end of each trial. Dashed lines mark the start of each stimulus, and solid lines mark the stimulus offset. n = 11 (no reward, low), 13 (water, low), 10 (sucrose, low), 10 (no reward, high), 8 (water, high), 10 (sucrose, high) mice. (D) Summary of the decisions across the first 10 trials under different threat and reward conditions. n = 11, 13, 10, 10, 8, 10 mice. (E–H) Escape distance under threat, duration in the reward zone, peak escape speed, and latency to flee across trials in all conditions. Shading denotes the standard error of the mean. (I) Transition trials marking the shift between phases across conditions. Red lines indicate the median; boxes span the interquartile range (IQR); whiskers extend to 1.5 × IQR beyond the box.

Earlier version: During the first trial, mice showed distinct trajectories and speed profiles across threat and reward conditions (Figures 2B and 2C). These patterns changed rapidly within the first 10 trials, indicating fast habituation to the looming stimulus. This learning was reflected in altered decision patterns (Figure 2D) and changes in escape distance, duration in the reward zone, peak escape speed, and latency to flee (Figures 2E–H). To account for both learning and individual variability, we segmented trials for each mouse into early and late phases using a change-point detection approach on the learning curves (Figures 2I, S3C, and S3D, see Methods and Materials).

Revised version: During the first trial, mice showed distinct trajectories and speed profiles across threat and reward conditions (Figures 2B–C). These patterns changed rapidly within the first 10 trials, indicating fast habituation to the looming stimulus. This learning was reflected in altered decision patterns (Figure 2D) and changes in escape distance, duration in the reward zone, peak escape speed, and latency to flee (Figures 2E–H). To account for both learning and individual variability, we segmented trials for each mouse into early and late phases using a change-point detection approach on the learning curves (Figures 2I and S3C–D, see Materials and methods).

Earlier version: In the early phase, behavior was shaped predominantly by the level of threat. Under higher threat, mice were more likely to choose direct escape behaviors (Figure 3A), with longer escape distances (Figure 3B) and less time spent in the reward zone (Figure 3C), suggesting a trade-off between threat avoidance and reward pursuit. Notably, latency to flee decreased with higher threat (Figure 3D), indicating heightened vigilance (Buck, 1966). This threat-dependent change in vigilance was further supported by the longer interval before re-entering the reward zone and slower foraging speed (Figures 3E, 3F, and S4), aligning with findings in birds that elevated predation risk increases vigilance and reduces feeding time (Caraco et al., 1980). In addition, escape speed increased significantly with threat (Figure 3G), reflecting the combined effects of perceived higher risk and heightened vigilance. The analyses restricted to the first trial yielded consistent results (Figures S5).

Revised version: In the early phase, behavior was shaped predominantly by the level of threat. Under higher threat, mice were more likely to choose direct escape behaviors (Figure 3A), with longer escape distances (Figure 3B) and less time spent in the reward zone (Figure 3C), suggesting a trade-off between threat avoidance and reward pursuit. Notably, latency to flee decreased with higher threat (Figure 3D), indicating heightened vigilance (Buck, 1966). This threat-dependent change in vigilance was further supported by the longer interval before re-entering the trigger zone and slower foraging speed (Figures 3E–F and S4), aligning with findings in birds that elevated predation risk increases vigilance and reduces feeding time (Caraco et al., 1980). In addition, escape speed increased significantly with threat (Figure 3G), reflecting the combined effects of perceived higher risk and heightened vigilance. Analyses restricted to the first trial yielded consistent results (Figure S5).

Mice make economic decisions modulated by vigilance.

Earlier version: (A) Distribution of decisions in the early phase for six experimental conditions. N = 43 trials from 11 mice (none, low), 42 trials from 13 mice (water, low), 29 trials from 10 mice (sucrose, low), 31 trials from 10 mice (none, high), 21 trials from 8 mice (water, high) and 32 trials from 10 mice (sucrose, high), Chi-squared test. (B-G) Escape distance under threat, duration in the reward zone, latency to flee, foraging interval, foraging speed, and peak escape speed in the early phase across all conditions. Scheirer–Ray–Hare test with post hoc Dunn’s test (Holm correction). N = 43, 42, 29, 31, 21, 32 trials for B, C, F, G; N = 35, 37, 24, 31, 21, 32 trials for D; N = 116, 84, 54, 58, 36, 48 intervals for E. (H) Distribution of decisions in the late phase for six experimental conditions. N = 59 trials from 11 mice (none, low), 88 trials from 13 mice (water, low), 71 trials from 10 mice (sucrose, low), 67 trials from 10 mice (none, high), 59 trials from 8 mice (water, high) and 48 trials from 9 mice (sucrose, high), Chi-squared test. (I-N) Escape distance under threat, duration in the reward zone, latency to flee, foraging interval, foraging speed, and peak escape speed in the early phase across all conditions. Scheirer–Ray–Hare test with post hoc Dunn’s test (Holm correction). N = 59, 88, 71, 67, 59, 48 trials for I, J, M, N; N = 43, 49, 15, 66, 59, 47 trials for K; N = 116, 182, 114, 137, 71, 57 intervals for L. For all panels: *p < 0.05, **p < 0.01, ***p < 0.001.

Revised version: (A) Distribution of decisions in the early phase under six experimental conditions. n = 43 trials from 11 mice (no reward, low), 42 trials from 13 mice (water, low), 29 trials from 10 mice (sucrose, low), 31 trials from 10 mice (no reward, high), 21 trials from 8 mice (water, high) and 32 trials from 10 mice (sucrose, high), Chi-squared test. (B–G) Escape distance under threat, duration in the reward zone, latency to flee, foraging interval, foraging speed, and peak escape speed in the early phase across all conditions. Scheirer–Ray–Hare test with post hoc Dunn’s test (Holm correction). n = 43, 42, 29, 31, 21, 32 trials for B, C, F, G; n = 35, 37, 24, 31, 21, 32 trials for D; n = 116, 84, 54, 58, 36, 48 intervals for E. (H) Distribution of decisions in the late phase under six experimental conditions. n = 59 trials from 11 mice (no reward, low), 88 trials from 13 mice (water, low), 71 trials from 10 mice (sucrose, low), 67 trials from 10 mice (no reward, high), 59 trials from 8 mice (water, high) and 48 trials from 9 mice (sucrose, high), Chi-squared test. (I-N) Escape distance under threat, duration in the reward zone, latency to flee, foraging interval, foraging speed, and peak escape speed in the late phase across all conditions. Scheirer–Ray–Hare test with post hoc Dunn’s test (Holm correction). n = 59, 88, 71, 67, 59, 48 trials for I, J, M, N; n = 43, 49, 15, 66, 59, 47 trials for K; n = 116, 182, 114, 137, 71, 57 intervals for L. For all panels: *p < 0.05, **p < 0.01, ***p < 0.001.

Earlier version: In the late phase, both threat level and reward value contributed to behavior. Threat remained a dominant factor, while reward exerted a significant, context-dependent influence. Under low-threat conditions, decisions were primarily driven by perceived reward value. As reward value increased, mice exhibited fewer defensive responses (Figure 3H), indicating that they weighed risk and reward to make economic decisions. Higher rewards led to shorter escape distance (Figure 3I) and longer stays in the reward zone (Figure 3J), supporting value-based decision-making strategies. In contrast, vigilance-related metrics, including latency to flee, foraging interval, and foraging speed, were largely unchanged across reward conditions (Figure 3K–M), suggesting that vigilance remains relatively stable. Consistently, escape speed decreased as reward value increased (Figure 3N). The value-based decision-making under low-threat conditions was further validated using within-subject comparisons to control for individual variability (Figure S6).

Revised version: In the late phase, both threat level and reward value contributed to behavior. Threat remained a dominant factor, while reward exerted a significant, context-dependent influence. Under low-threat conditions, decisions were primarily driven by perceived reward value. As reward value increased, mice exhibited fewer defensive responses (Figure 3H), indicating that they weighed threat and reward to make economic decisions. Higher rewards led to shorter escape distance (Figure 3I) and longer stays in the reward zone (Figure 3J), supporting value-based decision-making strategies. In contrast, vigilance-related metrics, including latency to flee, foraging interval, and foraging speed, were largely unchanged across reward conditions (Figures 3K–M), suggesting that vigilance remained relatively stable. Consistently, escape speed decreased as reward value increased (Figure 3N). The value-based decision-making under low-threat conditions was further validated using within-subject comparisons to control for individual variability (Figure S6).

Earlier version: Under high-threat conditions, however, increasing reward value produced the opposite effect: mice showed more direct escape behaviors (Figure 3H). This shift in decision patterns can be attributed to heightened vigilance driven by reward, as evidenced by shorter latencies to flee, longer foraging intervals, and slower foraging speeds (Figures 3 K–M). This counterintuitive observation suggests that under high threat, reward value indirectly shapes decision-making by modulating vigilance. In contrast, duration in the reward zone did not increase with reward value (Figure 3J). Collectively, these findings reveal how innate decision-making in response to looming stimuli is shaped by the dynamic interplay between perceived threat intensity, reward value, and vigilance.

Revised version: Under high-threat conditions, however, increasing reward value produced the opposite effect: mice showed more direct escape behaviors (Figure 3H). This shift in decision patterns was accompanied by behavioral signatures of heightened vigilance, including shorter latencies to flee, longer foraging intervals, and slower foraging speeds (Figures 3K–M). These results point to a counterintuitive mechanism: under high threat, reward value does not directly promote continued foraging but instead shapes decisions by enhancing vigilance, which in turn biases animals toward rapid escape responses. Consistent with this interpretation, the duration spent in the reward zone did not increase with reward value (Figure 3J). Collectively, these findings reveal how innate decision-making in response to looming stimuli is shaped by the dynamic interplay between perceived threat intensity and reward value, with vigilance acting as a key modulating factor.

Influence of social hierarchy on decision-making under threat

Earlier version: Mice are social animals that live in groups, where social hierarchy often plays a critical role in shaping behavior. To investigate how social rank influences decision-making under threat, we compared the responses of dominant and subordinate mice to looming stimuli. Rank within each mouse pair was determined using the tube test before and after the behavioral experiments (Figure 4A, see Methods and Materials), and only pairs with stable rankings were included in the analysis.

Revised version: Mice are social animals that live in groups, where social hierarchy often plays a critical role in shaping behavior. To investigate how social rank influences decision-making under threat, we compared the responses of dominant and subordinate mice to looming stimuli. Rank within each mouse pair was determined using the tube test before and after the behavioral experiments (Figure 4A, see Materials and methods), and only pairs with consistent rankings were included in the analysis.

Influence of social hierarchy on innate decision-making.

Earlier version: (A) Schematic of the behavioral assay for studying the social modulation of innate decision-making. Top, experimental timeline. Each session lasted 2 hours per mouse pair during the pre-looming, looming, and post-looming phases. Bottom, schematic of the tube test. (B) Arena occupancy for an example pair of mice during the three sessions. (C) Visit frequency and total visit duration for dominant and subordinate mice during the pre-threat, threat, and post-threat sessions. N = 5 pairs (pre), 5 pairs (threat), 4 pairs (post), respectively; paired t-test. (D) Proportion of nest-return trials in escape trials for dominant and subordinate mice. N = 9 pairs, paired t-test. (E) Distance to the safe zone over seven days for an example pair. Looming stimuli were presented on days 4 and 5. (F) Percentage of time spent in the reward zone across days. Error bars represent SEM. N = 9 pairs, paired t-test. (G) Distance to the safe zone (left) and locomotion speed (right) for an example pair. (H) Behavioral decisions across the first 10 trials for 9 mouse pairs. (I) Pie charts showing decision distributions for dominant and subordinate mice. N = 90, 90 trials, Stuart-Maxwell test. (J) Latency to flee for dominant and subordinate mice. N = 78 trials; paired t test. (K) Foraging interval for dominant and subordinate mice. N = 75 trials. (L) Foraging speed for dominant and subordinate mice. N = 53 trials. (M–O) Escape distance under threat, peak escape speed, and duration in the reward zone for dominant and subordinate mice. N = 90 trials. *p < 0.05, **p < 0.01, ***p < 0.001.

Revised version: (A) Schematic of the behavioral assay for studying the social modulation of innate decision-making. Top, experimental timeline. Each session lasted 2 hours per mouse pair during the pre-threat, threat, and post-threat phases. Bottom, schematic of the tube test. (B) Arena occupancy for an example pair of mice during the three sessions. (C) Visit frequency and total visit duration for dominant and subordinate mice during the pre-threat, threat, and post-threat sessions. n = 5 pairs (pre), 5 pairs (threat), 4 pairs (post), respectively; paired t-test. (D) Proportion of nest-return trials in escape trials for dominant and subordinate mice. n = 9 pairs, paired t-test. (E) Distance to the safe zone over seven days for an example pair. Looming stimuli were presented on days 4 and 5. (F) Percentage of time spent in the reward zone across days. Error bars represent SEM. n = 9 pairs, paired t-test. (G) Distance to the safe zone (left) and locomotion speed (right) for an example pair. (H) Behavioral decisions across the first 10 trials for 9 mouse pairs. (I) Pie charts showing decision distributions for dominant and subordinate mice. n = 90, 90 trials, Stuart-Maxwell test. (J) Latency to flee for dominant and subordinate mice. n = 78 trials; paired t test. (K) Foraging interval for dominant and subordinate mice. n = 75 trials. (L) Foraging speed for dominant and subordinate mice. n = 53 trials. (M–O) Escape distance under threat, peak escape speed, and duration in the reward zone for dominant and subordinate mice. n = 90 trials. *p < 0.05, **p < 0.01, ***p < 0.001.

Earlier version: We first verified that behavioral responses to the visual threat were not confounded by the presence of a social partner. Before looming exposure, subordinate mice spent more time exploring the arena than dominant mice (Figures 4B and 4C), which may reflect social avoidance. However, this difference disappeared during and after looming exposure, suggesting that the looming stimulus became the primary driver of behavior. To confirm this, we examined the probability of mice fleeing to the nest. If subordinates were avoiding dominants, they should preferentially escape to the safe zone rather than return to the nest where the dominant was present. Contrary to this prediction, dominant and subordinate mice showed similar probabilities of fleeing directly to the nest (Figure 4D), confirming that defensive responses were driven by the visual threat rather than by social avoidance.

Revised version: We first verified that behavioral responses to the visual threat were not confounded by the presence of a social partner. Before looming exposure, subordinate mice spent more time exploring the arena than dominant mice (Figures 4B–C), which may reflect social avoidance. However, this difference disappeared during and after looming exposure, suggesting that the looming stimulus became the primary driver of behavior. To confirm this, we examined the probability that mice fled directly to the nest. If subordinates were avoiding dominants, they should preferentially escape to the safe zone rather than return to the nest occupied by the dominant mouse. Contrary to this prediction, dominant and subordinate mice showed similar probabilities of fleeing directly to the nest (Figure 4D), confirming that defensive responses were driven by visual threat rather than social avoidance.

Earlier version: When the threat co-localized with the reward, only dominant mice reduced their relative time spent in the reward zone during the 2-h session (Figures 4E and 4F), indicating greater vigilance and risk aversion compared to subordinates. This heightened vigilance in dominant mice was further supported by reduced habituation to the looming threat (Figures 4G, 4H, and S7), a higher proportion of direct escape decisions (Figure 4I), and shorter latencies to flee (Figure 4J). No rank differences were observed in foraging interval or foraging speed. This may be due to the experimental design, in which only one mouse was allowed in the arena at a time, potentially altering natural foraging dynamics.

Revised version: When the threat co-localized with the reward, only dominant mice reduced their relative time spent in the reward zone during the 2-h session (Figures 4E–F), indicating greater vigilance and risk aversion compared to subordinates. This heightened vigilance in dominant mice was further supported by reduced habituation to the looming threat (Figures 4G–H and S7), a higher proportion of direct escape decisions (Figure 4I), and shorter latencies to flee (Figure 4J). No rank differences were observed in foraging intervals or foraging speeds (Figures 4K–L). This may stem from the experimental design, in which only one mouse was allowed in the arena at a time, potentially limiting the sensitivity of these measures to detect rank-dependent differences in vigilance.

Earlier version: Furthermore, dominant mice prioritized threat avoidance over reward, fleeing longer distances with higher escape speed and spending less time in the reward zone (Figures 4M–O). To rule out the possibility that the tube test itself influenced defensive behavior, we conducted additional experiments in which looming sessions occurred both before and after the tube test (Figure S8); the results remained largely consistent. Together, these findings indicate that social hierarchy biases the trade-off between reward seeking and threat avoidance, with dominant mice exhibiting heightened threat vigilance and subordinate mice showing stronger reward-driven behavior.

Revised version: Furthermore, dominant mice prioritized threat avoidance over reward, fleeing longer distances at higher escape speeds and spending less time in the reward zone (Figures 4M–O). To rule out the possibility that the tube test itself influenced defensive behavior, we conducted additional experiments in which looming sessions occurred both before and after the tube test (Figure S8); the results remained largely consistent. Together, these findings indicate that social hierarchy biases the trade-off between reward seeking and threat avoidance, with dominant mice exhibiting heightened threat vigilance and subordinate mice showing stronger reward-driven behavior.

A mathematical model for innate decision-making under threat

Earlier version: Our experimental results demonstrate that innate decision-making in response to visual threats is influenced by perceived threat intensity, reward value, and vigilance. To understand the underlying decision-making process, we developed a mathematical model in which the evidence for escape is accumulated by a drift-diffusion leaky integrator. An escape response is triggered when the evidence level crosses a predefined threshold:

Revised version: Our experimental results demonstrate that innate decision-making in response to visual threats is influenced by perceived threat intensity, reward value, and vigilance. To understand the underlying decision-making process, we developed a mathematical model in which the evidence for escape is accumulated by a drift-diffusion leaky integrator. An escape response is triggered when the evidence level crosses a predefined threshold: (1) dxidt=−α⋅xi(t)+β⋅s(t)−r+δ⋅dWdt

Earlier version: Here, xi(t) is the accumulated escape evidence at time t on the ith trial. The leakage rate α drives the evidence towards zero and is the reciprocal of the integration time constant. The stimulus function s(t) denotes the normalized diameter of the looming stimulus (Figure S9A). Together with the threat gain β, which reflects stimulus contrast and the animal’s vigilance level, these terms represent the perceived threat level and drive evidence accumulation towards escape. The term r represents the perceived reward value, which suppresses the accumulation of escape evidence. W (t) is a Wiener process, such that W (t + Δt) − W (t) ∼ N (0, Δt), and δ is the diffusion rate. The binary variable E(t) indicates whether an escape decision is made, with H denoting the Heaviside step function and xthr the decision threshold. If a decision is reached within the first expansion of the looming disc, it is classified as a direct escape; otherwise, it is classified as an escape after assessment.

Revised version: Here, xi(t) is the accumulated escape evidence at time t on the ith trial. The leakage rate α drives the evidence toward zero and is the reciprocal of the integration time constant. The stimulus function s(t) denotes the normalized diameter of the looming stimulus (Figure S9A). Together with the threat gain β, which reflects both stimulus contrast and the animal’s vigilance level, these terms represent the perceived threat level and drive evidence accumulation toward escape. The term r represents the perceived reward value, which suppresses the accumulation of escape evidence. W (t) is a Wiener process, such that W(t+Δt)−W(t)∼N(0,Δt), and δ is the diffusion rate. The binary variable E(t) indicates whether an escape decision is made, with H denoting the Heaviside step function and xthr the decision threshold. If a decision is reached within the first expansion of the looming disc, it is classified as a direct escape; otherwise, it is classified as an escape after assessment.

Earlier version: Model parameters were estimated using behavioral data in the late phase of reward-threat experiments in two steps (see Methods and Materials). First, we fit the model to the data under no-reward conditions and identified the optimal leakage rate (α = 1.78), threat gain (low threat: β = 1; high threat: β = 1.44), diffusion rate (δ = 5.6), and decision threshold (xthr = 0.71) (Figure S9B). This revealed distinct decision-making dynamics under low- and high-threat conditions. Under low threat, the average evidence level remained below the escape threshold, and escape decisions emerged from stochastic fluctuations in individual trials (Figure 5A). Under high threat, the average evidence level crossed the threshold, indicating that escape decisions were mainly driven by threat gain (Figure 5B). Notably, threat gain and diffusion rate had distinct effects on latency to flee: threat gain shifted the mean latency, while diffusion rate affected its variance. Thus, the model captured not only the observed decision patterns but also the distributions of latency to flee.

Revised version: Model parameters are estimated using behavioral data from the late phase of the reward–threat experiments in two steps (see Materials and methods). First, we fit the model to the data under no-reward conditions to obtain the optimal leakage rate (α = 1.78), threat gain (low threat: β = 1; high threat: β = 1.44), diffusion rate (δ = 5.6), and decision threshold (xthr= 0.71) (Figure S9B). These parameters reveal distinct decision-making dynamics under low- and high-threat conditions. Under low threat, the average evidence level remains below the escape threshold, and escape decisions emerge from stochastic fluctuations in individual trials (Figure 5A). Under high threat, the average evidence level crosses the threshold, indicating that escape decisions are mainly driven by threat gain (Figure 5B). Notably, threat gain and diffusion rate exert distinct effects on latency to flee: threat gain shifts the mean latency, while diffusion rate affects its variance. Thus, the model captures not only the observed decision patterns but also the distributions of latency to flee.

Drift-diffusion leaky integrator model for escape decisions.

Earlier version: (A-F) Simulated accumulation of escape evidence, along with predicted latencies to flee and decision distributions across six threat and reward conditions. Green dashed horizontal lines mark the x = 0; blue dashed horizontal lines mark the x = xthr; red dashed vertical lines mark the onset of each looming stimulus; blue vertical lines mark the median. (G) Heatmap of decision scores as a function of threat gain and reward value. White and red stars indicate fitted parameters for the early and late phases of the reward–threat paradigm, respectively; upward and downward pink triangles indicate fitted parameters for dominant and subordinate mice in the social-threat paradigm, respectively. (H) Heatmap of latencies to flee as a function of threat gain and reward value. (I) Vigilance as a function of input strength, illustrating how the indirect effect of reward on defensive decisions via vigilance depends on the baseline vigilance level. (J) Schematic showing how threat intensity, reward value, and vigilance jointly determine defensive decisions. Color saturation indicates the likelihood of defensive decisions.

Revised version: (A-F) Simulated accumulation of escape evidence, along with predicted latencies to flee and decision distributions across six threat and reward conditions. Green dashed horizontal lines mark the x = 0; blue dashed horizontal lines mark the x = xthr; red dashed vertical lines mark the onset of each looming stimulus; blue vertical lines mark the median. (G) Heatmap of decision scores as a function of threat gain and reward value. White and red stars indicate fitted parameters for the early and late phases of the reward–threat paradigm, respectively; upward and downward pink triangles indicate fitted parameters for dominant and subordinate mice in the social-threat paradigm, respectively. (H) Heatmap of latencies to flee as a function of threat gain and reward value. (I) Vigilance as a function of input strength, illustrating how the indirect effect of reward on defensive decisions via vigilance depends on the baseline vigilance level. (J) Schematic showing how threat intensity, reward value, and vigilance jointly determine defensive decisions. Color saturation indicates defense likelihood.

Earlier version: In the second step, we identified the threat gain and reward value across different threat and reward conditions in the late phase (Figures S9C and S9D). Consistent with experimental findings that decisions under low threat were primarily driven by reward value, the fitted reward-value parameter increases with reward magnitude (0, 0.08, and 0.25 for no reward, water, and sucrose, respectively). With these parameters, the model reproduced the observed reduction in escape likelihood without changes in latency to flee (Figures 5A, 5C, and 5E). The same reward-value parameters were then used to model the decision-making process under high-threat conditions. The fitted threat-gain parameter increased across reward conditions (1.44, 1.8, and 2.55 for no reward, water, and sucrose, respectively), counteracting—and ultimately reversing—the influence of reward value. The model again reproduced the experimentally observed increase in escape decisions and decrease in latency to flee (Figures 5B, 5D, and 5F).

Revised version: In the second step, we identify threat gain and reward value across different threat and reward conditions in the late phase (Figures S9C–D). Consistent with experimental findings that decisions under low threat are primarily driven by reward value, the fitted reward-value parameter increases with reward magnitude (0, 0.08, and 0.25 for no reward, water, and sucrose, respectively). With these parameters, the model reproduces the observed reduction in escape likelihood without changes in latency to flee (Figures 5A, 5C, and 5E). The same reward-value parameters are then used to model the decision-making process under high-threat conditions. Here, the fitted threat-gain parameter increases across reward conditions (1.44, 1.8, and 2.55 for no reward, water, and sucrose, respectively), counteracting—and ultimately reversing—the influence of reward value. The model again reproduces the experimentally observed increase in escape decisions and decrease in latency to flee (Figures 5B, 5D, and 5F).

Earlier version: To assess the model’s predictive power, we simulated how the decision score and latency to flee vary as functions of reward value (r) and threat gain (β). Increasing reward generally reduces escape probability and increases latency to flee, whereas increasing threat gain has the opposite effect (Figures 5G, 5H, and S9E–G). These relationships were confirmed by a simplified deterministic model (Figures S9H–K, see Methods and Materials). Notably, reward can also indirectly promote escape by elevating vigilance, with the magnitude of this effect depending on the level of baseline vigilance, defined as vigilance in the absence of reward (Figure 5I). This dual influence explains the absence of a reward effect in the early phase under both threat conditions: high vigilance shifts the operating range to a region where reward’s indirect promoting effect is balanced by its direct suppressing effect on escape decisions.

Revised version: To assess the model’s predictive power, we simulate how decision score (see Materials and methods) and latency to flee vary as functions of reward value (r) and threat gain (β). Increasing reward generally reduces decision score and increases latency to flee, whereas increasing threat gain has the opposite effect (Figures 5G–H, and S9E–G). These relationships are confirmed by a simplified deterministic model (Figures S9H–K, see Materials and methods). Notably, reward can also indirectly promote escape by elevating vigilance, with the magnitude of this effect depending on the level of baseline vigilance, defined as vigilance in the absence of reward (Figure 5I). This dual influence explains the absence of a reward effect in the early phase under both threat conditions: high vigilance shifts the operating range to a region where reward’s indirect promoting effect is balanced by its direct suppressing effect on escape decisions.

Earlier version: We next applied this model to the social-threat paradigm. The behavior of dominant mice was captured by a combination of lower reward value and higher threat gain compared to subordinates (Figures 5G and 5H). This aligns with the experimental finding that dominant animals spent less time in the reward zone and fled more quickly (Figures 4J and 4O). Together, these results demonstrate that the model robustly captures key behavioral features across phases and paradigms, providing a unified mechanistic account of innate decision-making under threat (Figure 5J).

Revised version: We next apply this model to the social-threat paradigm. The behavior of dominant mice is captured by a combination of lower reward value and higher threat gain compared to subordinates (pink triangles in Figures 5G–H). This aligns with the experimental finding that dominant animals spent less time in the reward zone and fled more quickly (Figures 4J and 4O). Together, these results demonstrate that the model robustly captures key behavioral features across phases and paradigms, providing a unified mechanistic account of innate decision-making under threat (Figure 5J).

Discussion

Earlier version: When confronted by a predator, an animal must rapidly decide whether to escape, freeze, or continue its ongoing behavior. This life-or-death decision is crucial for rodents facing aerial predators. While mice are social animals and most predator encounters occur during foraging, it remains unclear how they integrate perceived risk and reward when making defensive decisions, or how these decisions are influenced by social hierarchy. Using a foraging-based paradigm, we demonstrate that responses to visual threat are shaped by threat intensity, reward value, social rank, and internal state, and this integrated decision process can be captured by a mathematical model.

Revised version: When confronted by a predator, an animal must rapidly decide whether to escape, freeze, or continue its ongoing behavior. This life-or-death decision is especially critical for rodents facing aerial predators. Although mice are social animals and most predator encounters occur during foraging, it remains unclear how they integrate perceived threat and reward when making defensive decisions, or how these decisions are further influenced by social hierarchy. Using a foraging-based paradigm, we demonstrate that responses to visual threat are shaped by threat intensity, reward value, social rank, and internal state, and that this integrated decision process can be captured by a drift-diffusion leaky integrator.

Main findings

Earlier version: We designed a behavioral paradigm to simulate how mice respond to visual threats during foraging (Figure 1). Mice rapidly adapted to repeated looming threats, with substantial inter-individual variability in the rate of habituation (Figure 2). While threat intensity robustly shaped behavior in both early and late phases of habituation, reward exerted a clear influence only in the late phase (Figure 3). Specifically, increasing threat intensity shifted behavior towards more defensive responses. In contrast, reward value produced a context-dependent effect: under low-threat conditions, higher rewards suppressed defensive responses, aligning with value-based decision theory; whereas under high-threat conditions, higher rewards promoted escape decisions, potentially reflecting heightened vigilance. Innate decision-making was further shaped by social hierarchy (Figure 4): dominant mice exhibited a stronger tendency towards defensive behaviors, whereas subordinates were reward-driven and less likely to flee. Finally, we developed a drift-diffusion leaky integrator model to simulate the decision-making process, effectively capturing the key features of the observed decision patterns (Figure 5).

Revised version: We designed a behavioral paradigm to simulate how mice respond to visual threats during foraging (Figure 1). Mice adapted rapidly to repeated looming stimuli, showing substantial inter-individual variability in the rate of habituation (Figure 2). While threat intensity robustly shaped behavior in both early and late phases of habituation, reward exerted a clear influence only in the late phase (Figure 3). Specifically, increasing threat intensity shifted behavior toward more defensive responses. Reward value, in contrast, produced a context-dependent effect: under low-threat conditions, higher rewards suppressed defensive responses, consistent with value-based decision theory; under high-threat conditions, higher rewards promoted escape decisions, likely reflecting heightened vigilance. Innate decision-making was further shaped by social hierarchy (Figure 4): dominant mice exhibited a stronger bias toward defensive behaviors, whereas subordinates were more reward-driven and less likely to flee. Finally, we developed a drift-diffusion leaky integrator model that captures the key features of the observed behavioral patterns, providing a unified framework for understanding how threat, reward, and social context interact to shape innate defensive decisions (Figure 5).

Relation to earlier work

Earlier version: Aerial predators pose a significant threat to rodents. Previous studies have shown that overhead visual stimuli in the laboratory elicit robust defensive behaviors in rodents (Wallace et al., 2013; Yilmaz and Meister, 2013). Because these stimuli can mimic distinct predatory actions, the resulting behavioral responses depend on the specific physical properties of the stimulus. For example, an overhead expanding dark disc that mimics an approaching aerial predator preferentially triggers escape rather than freezing (Yilmaz and Meister, 2013), while a small black moving dot that mimics a cruising predator predominantly induces freezing behavior (De Franceschi et al., 2016). These findings suggest that the response is not a simple reflex but involves a decision-making process. Additional stimulus properties that influence the action selection include contrast, speed, size, and shape (De Franceschi et al., 2016; Evans et al., 2018; Yang et al., 2020). Support for a decision-making process also comes from its modulation by environmental factors and experience: mice freeze more when no refuge is available (Wei et al., 2015) and quickly learn that the looming stimulus poses no real threat (Lenzi et al., 2022; Vale et al., 2017; Zhong et al., 2023). This innate decision-making process is observed across invertebrates and vertebrates (Evans et al., 2019).

Revised version: Aerial predators pose a significant threat to rodents. Previous studies have shown that overhead visual stimuli in the laboratory elicit robust defensive behaviors in rodents (Wallace et al., 2013; Yilmaz and Meister, 2013). Because these stimuli can mimic distinct predatory actions, the resulting behavioral responses depend on the specific physical properties of the stimulus. For example, an overhead expanding dark disc that mimics an approaching aerial predator preferentially triggers escape rather than freezing (Yilmaz and Meister, 2013), while a small black moving dot that mimics a cruising predator predominantly induces freezing behavior (De Franceschi et al., 2016). These findings suggest that the response is not a simple reflex but involves a decision-making process. Additional stimulus properties that influence action selection include contrast, speed, size, and shape (De Franceschi et al., 2016; Evans et al., 2018; Yang et al., 2020). Support for a decision-making process also comes from its modulation by environmental factors and experience: mice freeze more when no refuge is available (Wei et al., 2015) and quickly learn that the looming stimulus poses no real threat (Lenzi et al., 2022; Vale et al., 2017; Zhong et al., 2023). This innate decision-making process is observed across invertebrates and vertebrates (Evans et al., 2019).

Earlier version: Compared with earlier work, the present study differs in two aspects. First, instead of manually annotating behavioral responses, we employed a machine-learning approach to classify behavioral decisions. This approach largely reduced the labor required for labeling and minimized misclassi-fications due to inter-individual variability. Second, unlike previous studies using 2D arenas, we designed a 1D linear arena to simulate foraging conditions, where the nest and reward zone were separated by a long corridor. Decision patterns observed here were similar to those reported in 2D environments (Yang et al., 2020). Consistent with previous findings (Lenzi et al., 2022), mice quickly adapted to the looming stimulus after a few trials.

Revised version: Methodologically, the present study differs from earlier work in two aspects. First, instead of manually annotating behavioral responses, we employed a machine-learning approach to classify behavioral decisions. This approach substantially reduced the labor required for labeling and minimized misclassifications due to inter-individual variability. Second, unlike previous studies using 2D arenas, we designed a 1D linear arena to simulate foraging conditions, where the nest and reward zone were separated by a long corridor. Importantly, the decision patterns observed in our paradigm are similar to those reported in 2D environments (Yang et al., 2020), and mice quickly adapted to the looming stimulus after a few trials, consistent with previous findings (Lenzi et al., 2022).

Earlier version: One concern with the linear arena design is that the looming stimulus is usually presented at the end of the arena, raising the possibility that mice simply retreat upon reaching a physical boundary. To address this, we presented the stimulus between the foraging mice and the nest (Figure S2A). Under this condition, mice quickly escape towards the nest—and thus towards the threat—rather than away from it (Figure S2B). This behavior indicates that their escape responses are not simple reflexes, but instead incorporate the safety of the nest into their decisions.

Revised version: One concern with the linear arena design was that the looming stimulus was usually presented at the end of the arena, raising the possibility that mice simply retreated upon reaching a physical boundary. To address this, we presented the stimulus between the foraging mice and the nest (Figure S2A). Under this condition, mice quickly fled toward the nest—and thus toward the threat—rather than away from it (Figure S2B, Video 5). This behavior indicates that their escape responses are not simple reflexes, but instead incorporate the safety of the nest into their decisions.

Earlier version: Fanselow and Lester’s predatory imminence continuum theory posits that defensive behaviors are graded according to perceived threat level (Fanselow and Lester, 1988). To test whether this frame-work applies to aerial threats, we varied both prey–threat distance and prey–safety distance. As prey-threat distance increased, mice showed less direct escape behavior, with longer latencies to flee and slower escape speeds (Figures S2C–G), indicating that defensive responses scale with threat proximity. These effects cannot be explained by the failure to detect the stimulus at longer distances; if that were the case, foraging speed would remain unaffected. Instead, mice slowed as they approached the threat in the 75-cm condition (Figure S2D), reflecting lower fear compared to shorter distances. In contrast, prey-safety distance did not significantly influence defensive behaviors (Figures S2H and S2I). One interpretation is that once an aerial threat is detected, the urgency to escape overrides evaluation of refuge proximity. To further verify this, we introduced barriers that lengthened the return path to the nest. Even under these conditions, defensive behaviors remained unaffected by prey-safety distance (Figures S2J and S2K).

Revised version: Fanselow and Lester’s predatory imminence continuum theory posits that defensive behaviors are graded according to perceived threat level (Fanselow and Lester, 1988). To test whether this frame-work applies to aerial threats, we varied both prey–threat distance and prey–safety distance. As prey–threat distance increased, mice showed less direct escape behavior, with longer latencies to flee and slower escape speeds (Figures S2C–G, Videos 6–8), indicating that defensive responses scale with threat proximity. These effects cannot be explained by the failure to detect the stimulus at longer distances; if that were the case, foraging speed would remain unaffected. Instead, mice slowed as they approached the threat in the 75-cm condition (Figure S2D). In contrast, prey–safety distance did not significantly influence defensive behaviors (Figures S2H–I). One interpretation is that once an aerial threat is detected, the urgency to escape overrides evaluation of refuge proximity. To further verify this, we introduced barriers that lengthened the return path to the nest. Even under these conditions, defensive behaviors remained unaffected by prey–safety distance (Figures S2J–K, Video 9).

Economic and social modulations of innate decision-making under threat

Earlier version: Value-based decision-making—where choices arise from comparisons of the subjective value of expected outcomes—is a well-established framework for understanding behavior across species (Rangel et al., 2008). Prospect theory, which accounts for decision-making under risk in humans (Kahneman and Tversky, 1979), has been applied to learned decision-making in rodents (Constantinople et al., 2019). Here, we demonstrate that value-based decision-making also extends to innate defensive behaviors: when confronted with threat, mice integrate perceived reward and threat value to make their decisions.

Revised version: Value-based decision-making—where choices arise from comparisons of the subjective value of ex-pected outcomes—is a well-established framework for understanding behavior across species (Rangel et al., 2008). Prospect theory, which accounts for decision-making under risk in humans (Kahneman and Tversky, 1979), has been applied to learned decision-making in rodents (Constantinople et al., 2019). Here, we demonstrate that value-based decision-making also extends to innate defensive behaviors: when confronted with a threat, mice integrate perceived reward and threat value to make their decisions.

Earlier version: Importantly, perceived value is determined not only by stimulus physical properties but also by internal brain state. In the resting state, for example, perceived stimulus strength scales with physical intensity according to Weber–Fechner laws. In our paradigm, the perceived threat value is strongly modulated by vigilance, a state of heightened and sustained attention to potential danger. This vigilance-dependent modulation substantially alters how the same looming stimulus is evaluated and the resulting decisions.

Revised version: Importantly, perceived value is determined not only by the physical properties of the stimulus but also by the animal’s internal brain state. In the resting state, for example, perceived stimulus strength scales with physical intensity according to the Weber–Fechner law. In our paradigm, the perceived threat value is strongly modulated by vigilance, a state of heightened and sustained attention to potential danger. This vigilance-dependent modulation substantially alters how the same looming stimulus is evaluated and how decisions are made.

Earlier version: Consistent with value-based evaluation, increasing threat intensity reliably increased defensive responses across reward conditions and behavioral phases. These changes reflect differences in perceived threat rather than stimulus detection, as evidenced by coordinated changes in escape latency, distance, and speed, and further supported by trial-by-trial inspection confirming reliable stimulus detection across contrast levels (Figure S5).

Revised version: Consistent with value-based evaluation, increasing threat intensity reliably increased defensive responses across reward conditions and behavioral phases. These changes reflect differences in perceived threat rather than stimulus detection, as evidenced by coordinated changes in escape latency, distance, and speed, and further supported by trial-by-trial inspection confirming reliable stimulus detection across contrast levels (Figure S3B).

Earlier version: Reward-dependent behavior also follows value-based principles, but in a vigilance-dependent manner. In nature, foraging often occurs under predation risk, and evolutionary pressure has favored prey that maintain heightened vigilance in the presence of high-value rewards. Consistent with this, increasing the reward magnitude increased both perceived reward value and—via elevated vigilance—perceived threat value. These opposing effects jointly govern the escape decision.

Revised version: Reward-dependent behavior also follows value-based principles, but in a vigilance-dependent manner. In nature, foraging often occurs under predation risk, and evolutionary pressure has favored prey that maintain heightened vigilance in the presence of high-value rewards. Consistent with this, increasing reward magnitude increased both perceived reward value and—via elevated vigilance—perceived threat value. These opposing effects jointly shaped the escape decision.

Earlier version: In the early phase, baseline vigilance is already high because looming is an innate threat. Under this high-vigilance condition, reward-driven increases in perceived reward and perceived threat counterbalance each other, resulting in no behavioral differences across reward conditions. With repeated exposure, habituation to looming reduces vigilance, shifting the operating region leftward along a sigmoidal vigilance–intensity function (Figure 5I). Such a sigmoidal form is biologically plausible because vigilance is bounded by a relaxed-state minimum and by cognitive capacity. This shift has opposite consequences under low versus high threat. Under low threat in the late phase, reduced vigilance compresses the dynamic range, weakening the indirect effect of reward on perceived threat and allowing perceived reward value to dominate decisions. Under high threat, the vigilance function remains within a steeper region of the sigmoid, where increases in reward effectively elevate vigilance and thus perceived threat. As a result, reward produces the opposite behavioral effect under high threat. Together, these dynamics explain why reward effects are absent early but emerge after habituation in a strongly threat-dependent manner.

Revised version: In the early phase, baseline vigilance is already high because looming represents an innate threat. Under this high-vigilance condition, increasing reward magnitude elevates both perceived reward and threat value; these opposing effects counterbalance each other, resulting in no behavioral differ-ences across reward conditions. With repeated exposure, habituation to looming reduces vigilance, shifting the operating region leftward along a sigmoidal vigilance-intensity function (Figure 5I). This sigmoidal form is biologically plausible, as vigilance is bounded by a relaxed-state minimum and by cognitive capacity limits. This shift has opposite consequences under low versus high threat. Under low threat in the late phase, reduced vigilance compresses the dynamic range, weakening the indirect effect of reward on perceived threat and allowing perceived reward value to dominate decisions. Under high threat, however, the vigilance function remains within a steeper region of the sigmoid, where increases in reward effectively elevate vigilance and thus perceived threat. As a result, reward produces the opposite behavioral effect under high threat. Together, these dynamics explain why reward effects are absent early but emerge after habituation in a threat-dependent manner.

Earlier version: Social rank–dependent differences in defensive behavior also fit within this framework. Even under identical conditions, dominant mice perceived higher threat value due to higher vigilance and lower reward value, consistent with their more risk-averse behavior. Similar status-dependent escape patterns have been reported across species, including voles (Kleiman et al., 2014) and crayfish (Krasne et al., 1997), suggesting an evolutionarily conserved strategy.

Revised version: Social rank-dependent differences in defensive behavior also fit within this framework. Even un-der identical conditions, dominant mice perceived a higher threat value—driven by their elevated vigilance—along with a lower reward value, consistent with their more risk-averse behavior. Similar status-dependent escape patterns have been reported across species, including voles (Kleiman et al., 2014) and crayfish (Krasne et al., 1997), suggesting an evolutionarily conserved strategy.

Earlier version: Because vigilance declines with habituation, factors that modulate the rate of habituation also influence defensive behavior. Although individuals vary, several patterns are consistent. Mice habituate more rapidly under high-threat or high-reward conditions (Figure 2I). Dominant mice habituate more slowly than subordinates, consistent with their elevated vigilance (Figures 5H and S7A). Housing conditions also alter habituation: group-housed animals adapt more slowly to looming than single-housed animals, and dominant–subordinate pairs habituate even more slowly than group-housed mice (Figures 2I and S7A). These patterns indicate that stable social relationships sustain high vigilance and slow habituation—an evolutionarily conserved strategy that may enhance survival.

Revised version: Because vigilance declines with habituation, factors that modulate the rate of habituation also influence defensive behavior. Although individuals vary, several patterns emerged. Mice habituated more rapidly under high-threat or high-reward conditions (Figure 2I). Dominant mice habituated more slowly than subordinates, consistent with their elevated vigilance (Figures 4H and S7A). Housing conditions also altered habituation: group-housed animals adapted more slowly to looming than single-housed animals, and dominant–subordinate pairs habituated even more slowly than group-housed mice (Figures 2I and S7A). These patterns indicate that stable social relationships sustain high vigilance and slow habituation—an evolutionarily conserved strategy that may enhance survival.

Earlier version: One limitation of the present study is that our experiments were conducted exclusively in male mice to avoid the confounding effects of behavioral variability associated with the female estrous cycle. Extending this paradigm to females—with careful control for estrous cycles—will provide a more comprehensive understanding of how economic and social factors modulate defensive decision-making.

Revised version: One limitation of the present study is that our experiments were conducted exclusively in male mice to avoid the confounding effects of behavioral variability associated with the female estrous cycle. Extending this paradigm to females—with careful control for estrous cycle—will provide a more comprehensive understanding of how economic and social factors modulate defensive decision-making.

A proposed role for the superior colliculus in value integration

Earlier version: What neural circuits underlie the economic and social modulations of decision-making under threat? A growing body of research implicates the SC as a key node in looming-evoked defensive behaviors (Evans et al., 2018; Shang et al., 2015; Wei et al., 2015). In parallel, reward is encoded by dopaminergic neurons in the ventral tegmental area (VTA) (Cohen et al., 2012; Schultz, 1998) and serotonergic neurons in the dorsal raphe nucleus (Liu et al., 2014; Miyazaki et al., 2011), which project widely to regions such as ventral striatum (Schultz et al., 1992), orbitofrontal cortex (Tremblay and Schultz, 2000), cerebellum (Wagner et al., 2017). Furthermore, social status modulates behavior via circuits involving the medial prefrontal cortex (mPFC) (Kingsbury et al., 2019; Wang et al., 2011) and serotonergic signaling (Edwards and Kravitz, 1997; Raleigh et al., 1991; Yeh et al., 1997).

Revised version: What neural circuits underlie the economic and social modulations of decision-making under threat? A growing body of research implicates the SC as a key node in looming-evoked defensive behav-iors (Evans et al., 2018; Shang et al., 2015; Wei et al., 2015). In parallel, reward is encoded by dopaminergic neurons in the ventral tegmental area (VTA) (Cohen et al., 2012; Schultz, 1998) and serotonergic neurons in the dorsal raphe nucleus (Liu et al., 2014; Miyazaki et al., 2011), which project widely to regions such as the ventral striatum (Schultz et al., 1992), orbitofrontal cortex (Trem-blay and Schultz, 2000), and cerebellum (Wagner et al., 2017). Furthermore, social status modulates behavior via circuits involving the medial prefrontal cortex (mPFC) (Kingsbury et al., 2019; Wang et al., 2011) and serotonergic signaling (Edwards and Kravitz, 1997; Raleigh et al., 1991; Yeh et al., 1997).

Earlier version: We propose the SC as a central hub for mediating the economic and social modulations of decision-making for several reasons. First, while neurons in the superficial (sSC) and intermediate (iSC) layers of the medial SC are involved in detecting looming threats and triggering defensive responses (Li et al., 2023, 2020), they also express receptors for both dopamine and serotonin (Mooney et al., 1996; Woolrych et al., 2021), indicating they can receive reward- and social-related information. Notably, dopamine receptor expression is layer-specific, with D1 receptors enriched in the sSC and D2 receptors in the iSC, suggesting layer-dependent processing of reward signals. Consistently, reward-related signals have been observed in the medial sSC (Baruchin et al., 2023). Second, the deep layers of SC receive from the retrosplenial cortex (RSC) (Campagner et al., 2023) and indirectly from the hippocampus (Benavidez et al., 2021), which may convey spatial and contextual information relevant to reward (Calvin et al., 2025). In parallel, the lateral SC mediates approach behaviors towards rewarding stimuli such as food or prey (Comoli et al., 2012; Krauzlis et al., 2013). These findings suggest that medial–lateral interactions within the SC may play a critical role in resolving threat–reward trade-offs during decision-making. Third, the deep-layer SC receives inputs from the mPFC (Benavidez et al., 2021), providing a pathway through which social and contextual signals can modulate defensive decisions. Finally, these SC deep layers are strongly and reciprocally connected with the periaqueductal gray (PAG) and project to dopaminergic neurons in the substantia nigra (Comoli et al., 2003), enabling efficient translation of integrated sensory, economic, and social signals into action execution. They also project to the lateral amygdala via the lateral posterior nucleus (LP) (Wei et al., 2015) to rapidly elicit the fear state.

Revised version: We propose the SC as a central hub for mediating the economic and social modulations of decision-making for several reasons. First, while neurons in the superficial (sSC) and intermediate (iSC) layers of the medial SC are involved in detecting looming threats and triggering defensive responses (Li et al., 2023, 2020), they also express receptors for both dopamine and serotonin (Mooney et al., 1996; Woolrych et al., 2021), indicating they can receive reward- and social-related information. Notably, dopamine receptor expression is layer-specific, with D1 receptors enriched in the sSC and D2 receptors in the iSC, suggesting layer-dependent processing of reward signals. Consistently, reward-related signals have been observed in the medial sSC (Baruchin et al., 2023). Second, the deep layers of SC receive inputs from the retrosplenial cortex (RSC) (Campagner et al., 2023) and indirectly from the hippocampus (Benavidez et al., 2021), which may convey spatial and contextual information relevant to reward (Calvin et al., 2025). In parallel, the lateral SC mediates approach behaviors toward rewarding stimuli, such as food or prey (Comoli et al., 2012; Krauzlis et al., 2013). These findings suggest that medial–lateral interactions within the SC may play a critical role in resolving threat–reward trade-offs during decision-making. Third, the deep-layer SC receives inputs from the mPFC (Benavidez et al., 2021), providing a pathway through which social and contextual signals can modulate defensive decisions. Finally, these SC deep layers are strongly and reciprocally connected with the periaqueductal gray (PAG) and project to dopaminergic neurons in the substantia nigra (Comoli et al., 2003), enabling efficient translation of integrated sensory, economic, and social signals into action execution. They also project to the lateral amygdala via the lateral posterior nucleus (LP) (Wei et al., 2015) to rapidly elicit the fear state.

Added in the revised version: While the anatomical and functional evidence support the SC as a central hub, we emphasize that the neural architecture underlying such decisions is far more complex. The precise integration of internal state, social rank, and external threat likely emerges from the dynamic interplay of distributed circuits, including prefrontal, limbic, and neuromodulatory systems. Dissecting how these circuits interact to shape defensive decisions remains an important direction for future research.

Removed in the revised version: Taken together, these anatomical and functional studies support the view that the SC serves as a central hub for the economic and social modulation of decision-making evoked by visual threat. In particular, cross-layer processing and medial–lateral interactions within the SC may provide key circuit mechanisms through which reward value and social context shape innate defensive decisions.

Mathematical modeling

Earlier version: The proposed drift-diffusion leaky integrator model builds on an integrator model of internal state (Gibson et al., 2015) and extends it by incorporating a reward-driven drift-diffusion process (Ratcliff, 1978). Conceptually, the integrator part in our model aligns with Lorenz’s “hydraulic” model of motivation, which describes how internal drives shape behavior (Lorenz, 1950). While our model resembles the leaky integrator used to model escape behavior in flies (Gibson et al., 2015), it differs in several ways. First, the earlier model lacks a reward-related drift component. Second, while they modeled the looming effect as a delta function, our model allows the evidence level to vary continuously with the stimulus size. Third, the influence of vigilance is absent in that model. A similar model has been proposed (Evans et al., 2018), but it did not incorporate reward and vigilance. Note that the evidence level for escape in our model is not equivalent to the fear level (Anderson and Adolphs, 2014); rather, when it crosses a threshold, the animal may enter a fear state.

Revised version: The proposed drift-diffusion leaky integrator model builds on an integrator model of internal state (Gib-son et al., 2015) and extends it by incorporating a reward-driven drift-diffusion process (Ratcliff, 1978). Conceptually, the integrator component in our model aligns with Lorenz’s “hydraulic” model of motivation, which describes how internal drives shape behavior (Lorenz, 1950). While our model resembles the leaky integrator used to model escape behavior in flies (Gibson et al., 2015), it differs in several ways. First, the earlier model lacks a reward-related drift component. Second, while that model treats the looming effect as a delta function, our model allows the evidence level to vary continuously with stimulus size. Third, the influence of vigilance is absent in that model. A similar model has been proposed (Evans et al., 2018), but it did not incorporate reward and vigilance. Note that the evidence level for escape in our model is not equivalent to the fear level (Anderson and Adolphs, 2014); rather, when it crosses a threshold, the animal may enter a fear state.

Earlier version: Below, we briefly discuss the roles of individual parameters. The leakage rate α represents a drive towards the resting state. Perceived threat was modeled as the product of the threat gain β and the sensory input s(t) and promotes escape decisions, whereas perceived reward r drives the decision away from escape. The threat gain increases not only with threat intensity but also with reward value via its influence on vigilance. The effect of reward on vigilance depends on the operating region of the sigmoidal vigilance function. This operating region is set by the baseline vigilance level (Figure 5I) and reflects habituation to repeated looming threats. Furthermore, threat gain and reward value have distinct effects on escape latency: higher threat gain exponentially shortened latency, whereas higher reward value increased it linearly (Figure S9F and S9G). These effects are confirmed using a simplified deterministic version of the model (Figure S9J and S9K; see Methods and Materials).

Revised version: Below, we briefly discuss the roles of individual parameters. The leakage rate α represents a drive toward the resting state. Perceived threat was modeled as the product of the threat gain β and the sensory input s(t) and promotes escape decisions, whereas perceived reward r drives the decision away from escape. The threat gain increases not only with threat intensity but also with reward magnitude via its influence on vigilance. The effect of reward on vigilance depends on the operating region of the sigmoidal vigilance function. This operating region is set by the baseline vigilance level (Figure 5I) and reflects habituation to repeated looming threats. Furthermore, threat gain and reward value have distinct effects on escape latency: higher threat gain exponentially shortened latency, whereas higher reward value increased it linearly (Figures S9F–G). These effects are confirmed using a simplified deterministic version of the model (Figure S9J–K; see Materials and methods).

Lastly, the diffusion rate δ plays a critical role, particularly under low-threat conditions where the average evidence trajectory fails to reach the threshold. This parameter may reflect the fluctuation of the animal’s internal state. The diffusion rate captures the trial-to-trial variability: even small amounts of noise can accumulate over time, resulting in large differences across individual trials and contributing to variability in escape latency. Overall, this computational framework not only accounts for our experimental observations but also offers a quantitative foundation for studying innate decision-making across species. Understanding how these parameters are implemented at the circuit level is an exciting direction for future research.

Earlier version: Methods and Materials

Revised version: Materials and methods

Added in the revised version: Animals

Removed in the revised version: Animal

Earlier version: Male C57BL/6J mice were group-housed on a 12-hour light / 12-hour dark cycle and used at ages 2-3 months. Behavioral experiments were conducted during the light phase. All experimental procedures were performed under the animal welfare guidelines and approved by the Institutional Animal Care and Use Committee at the Chinese Institute for Brain Research, Beijing.

Revised version: Male C57BL/6J mice were group-housed on a 12-hour light/12-hour dark cycle and used at 2–3 months of age. Behavioral experiments were conducted during the light phase. All experimental procedures were performed under the animal welfare guidelines and approved by the Institutional Animal Care and Use Committee at the Chinese Institute for Brain Research, Beijing.

Behavioral platform

Earlier version: The behavioral platform consisted of a nest, a linear arena, a radio frequency identification (RFID) system, a real-time mouse position detection system, and a reward delivery port. As illustrated in Figure 1A, the nest (40 cm (L) × 20 cm (W) × 30 cm (H)) and the linear arena (100 cm × 10 cm × 30 cm) were made with infrared transmitting acrylic to allow unobtrusive behavioral monitoring. They were connected by a one-way tunnel (20 cm × 3 cm × 3 cm) for entering the arena and a 20 cm × 5 cm × 30 cm safe zone with a 5 cm × 5 cm one-way door for returning to the nest. Because it wasn’t directly under the monitor, the safe zone was darker than the arena. Mice were labeled via implanted RFID tags, which were detected by an RFID reader placed around the tunnel. Only one mouse was allowed to enter the arena at a time. Specifically, when all mice were in the nest, the door between the nest and the tunnel was open, while the door between the tunnel and the arena remained closed. When a single mouse entered the tunnel, as detected by the RFID system, the door to the nest closed and the door to the arena opened, allowing the mouse to enter the arena. The mice were rewarded at the end of the arena with water or sucrose. Licking time and reward volume were recorded.

Revised version: The behavioral platform consisted of a nest, a linear arena, a radio frequency identification (RFID) system, a real-time mouse position detection system, and a reward delivery port. As illustrated in Figure 1A, the nest (40 (L) × 20 (W) × 30 (H) cm) and the linear arena (100 × 10 × 30 cm) were made of infrared-transmitting acrylic to allow unobtrusive behavioral monitoring. They were connected by a one-way tunnel (20 × 3 × 3 cm) for entering the arena and a 20 × 5 × 30 cm safe zone with a 5 × 5 cm one-way door for returning to the nest. Because it was not directly under the monitor, the safe zone was darker than the arena.

Added in the revised version: Mice were identified with implanted RFID tags, which were detected by an RFID reader positioned around the tunnel. Only one mouse was allowed to enter the arena at a time. Specifically, when all mice were in the nest, the door between the nest and the tunnel was open, while the door between the tunnel and the arena remained closed. When a single mouse entered the tunnel, as detected by the RFID system, the door to the nest closed and the door to the arena opened, allowing the mouse to enter the arena. Mice received rewards at the end of the arena in the form of water or 10% sucrose. Licking time and reward volume were recorded for each trial.

Earlier version: To track mouse position in real time, an OpenMV camera with a 90° field-of-view lens was mounted on the ground, 110 cm beneath the arena. The mouse position was used to control the tunnel doors and to trigger stimulus presentation when the mouse entered the arena. Specifically, the tunnel had two doors: the first connected the tunnel to the nest, and the second door connected the tunnel to the arena. Once a mouse entered the tunnel, the first door closed and the second door opened. A second camera (LBAS-U350-74M) with a lens (FA0615A) was also placed on the ground to record animal behavior at 30 frames per second.

Revised version: To track mouse position in real time, an OpenMV camera with a 90° field-of-view lens was mounted on the ground, 110 cm beneath the arena. The mouse position was used to control the tunnel doors and to trigger stimulus presentation when the mouse entered the arena. A second camera (LBAS-U350-74M) with a lens (FA0615A) was also placed on the ground to record animal behavior at 30 frames per second.

Visual stimulation

Earlier version: Looming stimuli were presented on a 55-inch monitor (121 cm × 68 cm) suspended 32 cm above the arena. Visual stimulation was implemented using Psychopy in Python and was aligned to the mouse’s real-time location. The stimulus was a black disc that expanded ten times on a gray background (∼ 65 cd / cm2). For each time, the disc expanded from 0° to 20° of visual angle at a speed of 40°/s, followed by a stationary phase lasting 0.3 s. The inter-stimulus interval was randomized between 1 and 2 minutes. Two stimulus contrasts at 20% and 99% were displayed by adjusting the luminance of the disc.

Revised version: Looming stimuli were presented on a 55-inch monitor (121 × 68 cm) suspended 32 cm above the arena. Visual stimulation was implemented using PsychoPy in Python and was aligned to the mouse’s real-time location. The stimulus consisted of a black disc that expanded ten times on a gray background (∼65 cd/cm2). On each expansion, the disc grew from 0° to 20° of visual angle at 40°/s, followed by a stationary phase lasting 0.3 s. The inter-stimulus interval was randomly varied between 1 and 2 minutes. Two stimulus contrasts (20% and 99%) were displayed by adjusting the luminance of the disc.

Economic modulation of innate decision-making

Earlier version: All mice were habituated to the behavioral platform for two days before the looming experiment. On the first day, five mice from the same home cage were placed in the nest for 30 minutes with all doors closed. Each mouse was then placed individually in the nest and allowed to explore the arena for 10 min under normal door operation; if a mouse entered the arena fewer than two times, the exploration period was extended until at least two arena visits were completed. Subsequently, all five mice were returned to the nest with all doors open and allowed to freely explore the arena for 2 h to ensure that each mouse learned the reward location at the end of the linear arena. On the second day, each mouse was placed individually in the nest and given an additional 1 h of exploration under normal door operation to further acclimate to the environment. The looming experiment was conducted the following day. An overhead looming stimulus of either low (20%) or high (99%) contrast was triggered when the mouse entered the trigger zone, defined as the region within 20 cm of the reward port. For each mouse, both the looming contrast and reward type remained constant across trials. In total, we collected and analyzed behavioral data from 590 trials across 62 mice.

Revised version: All mice were habituated to the behavioral platform for two days before the looming experiment. During both habituation and testing, three reward conditions were used: no reward, water, or 10% sucrose. Mice were not water-deprived in any of the groups. On the first day, five mice from the same home cage were placed in the nest for 30 minutes with all doors closed. Each mouse was then placed individually in the nest and allowed to explore the arena for 10 minutes under normal door operation; if a mouse entered the arena fewer than two times, the exploration period was extended until at least two arena visits were completed. Subsequently, all five mice were returned to the nest with all doors open and allowed to freely explore the arena for 2 hours to ensure that each mouse learned the reward location at the end of the linear arena. On the second day, each mouse was placed individually in the nest and given an additional 1 hour of exploration under normal door operation to further acclimate to the environment. The looming experiment was conducted the following day. An overhead looming stimulus of either low (20%) or high (99%) contrast was triggered when the mouse entered the trigger zone, defined as the region within 20 cm of the reward port. For each mouse, both the looming contrast and reward type remained constant across trials. In total, we collected and analyzed behavioral data from 590 trials across 62 mice.

Earlier version: To control for individual variation in decision-making, we compared the same animal’s behavior under reward and no-reward conditions. Two groups of mice were used in the experiment. In the first group (n = 5), mice were first tested under the reward condition, followed by the no-reward condition. In the second group (n = 4), the order was reversed. For the reward condition, mice were water-deprived one day before the exploration, and water was provided via the reward port. In the no-reward condition, mice were not water-deprived, and the reward port was removed. Before the looming experiment, mice were acclimated to the linear arena for two days as described above. On the third day, a looming stimulus with 20% contrast was displayed for five trials. In total, behavioral data from 84 trials across nine mice were recorded and analyzed.

Revised version: To control for individual variation in decision-making, we compared the same animal’s behavior under reward and no-reward conditions. Two groups of mice were used in this experiment. In the first group (n = 4), mice were first tested under the no-reward condition, followed by the reward condition. In the second group (n = 5), the order was reversed. For the reward condition, mice were water-deprived one day before exploration, and water was provided via the reward port. In the no-reward condition, mice were not water-deprived, and the reward port was removed. Before the looming experiment, mice were acclimated to the linear arena for two days as described above. On the third day, a looming stimulus with 20% contrast was displayed for five trials. In total, behavioral data from 84 trials across 9 mice were recorded and analyzed.

Social modulation of innate decision-making

Earlier version: To investigate the impact of social hierarchy on decision-making, we first determined the social rank of each mouse pair using the tube test (Wang et al., 2011). Mice were co-housed with a glass tube (3 cm diameter, 10 cm long) for one week, then trained to cross a 30 cm tube 10 times per day over two consecutive days. On the third day, a pair was tested for up to six trials; if one mouse achieved four consecutive wins, it was designated the dominant individual; otherwise, the pair (1 out of 6 pairs) was excluded from further experiment.

Revised version: To investigate the impact of social hierarchy on decision-making, we first determined the social rank of each mouse pair using the tube test (Lindzey et al., 1961; Wang et al., 2011), in which dominance is determined by the ability of one individual to force its opponent to retreat from a narrow tube. Mice were co-housed with a glass tube (3 cm diameter, 10 cm long) for one week, then trained to cross a 30-cm tube 10 times per day over two consecutive days. On the third day, each pair was tested for up to six trials; if one mouse achieved four consecutive wins, it was designated the dominant individual; otherwise, the pair (1 out of 6 pairs) was excluded from further experiments.

Earlier version: Before the looming experiment, mice were water-deprived and allowed to explore the linear arena for two hours per day over three days, during which water rewards were delivered at the end of the arena. Over the following two days, the looming experiment with a stimulus contrast of 99% was performed for two hours per day. After the looming experiment, mice continued to explore the arena for an additional two days, after which their social rank was reassessed with the tube test. All remaining pairs maintained a stable rank order.

Revised version: Before the looming experiment, mice were water-deprived and allowed to explore the linear arena for two hours per day over three days, during which water rewards were delivered at the end of the arena. Over the following two days, the looming experiment was conducted with a stimulus contrast of 99% for two hours per day. After the looming experiment, mice continued to explore the arena for an additional two days, after which their social rank was reassessed with the tube test. All remaining pairs maintained a stable rank order. In total, behavioral data from 180 trials across 18 mice were recorded and analyzed.

Earlier version: To assess whether the tube test itself influenced defensive decision-making, additional looming experiments were conducted on four mouse pairs before and after introducing the tube test. Specifically, the first looming experiment was conducted prior to the tube test in Figure 4A. The behavioral patterns in a total of 180 trials from 18 mice were recorded and analyzed.

Revised version: To assess whether the tube test itself influenced defensive decision-making, additional looming experiments were conducted on four mouse pairs before and after the tube test. Specifically, the first looming experiment was conducted prior to the tube test (Figure S8A). Behavioral data from 72 trials across 8 mice were recorded and analyzed.

Behavioral quantification

Earlier version: Animal behaviors were quantified in three steps. First, two key points—the nose and tail base—were labeled and tracked in the recorded videos using DeepLabCut (Figure S1A). Tracked points with likelihood scores below 0.6 were excluded and filled via linear interpolation. Locomotion speed was calculated for each frame and smoothed using a 0.3 s moving average window. Unless otherwise specified, the speed of the tail base was used for subsequent analyses. Second, individual 30 s trials were extracted, each consisting of 10 seconds before, 8 seconds during, and 12 seconds after the looming stimulus. Multiple episodes of these states could occur within a single trial.

Revised version: Animal behaviors were quantified in three steps. First, two key points—the nose and tail base—were labeled and tracked in the recorded videos using DeepLabCut (Mathis et al., 2018). Tracked points with likelihood scores below 0.6 were excluded and replaced using linear interpolation. Locomotion speed was calculated for each frame and smoothed using a 0.3 s moving average window. Unless otherwise specified, the speed of the tail base was used for subsequent analyses. Second, individual 30-second trials were extracted, each consisting of 10 s before, 8 s during, and 12 s after the looming stimulus.

Earlier version: Third, we defined 19 behavioral features related to locomotion speed, distance, and state transitions. Eleven features were related to speed and distance: (1) peak speed towards the reward port before stimulus onset; (2-4) maximum, mean, and coefficient of variation of speed after stimulus onset; (5) latency to peak speed; (6) hiding latency, defined as the time between stimulus onset and arrival at the safe zone (set to 20 s if the mouse did not reach the safe zone by the end of the trial); (7) escape distance, defined as the distance covered between stimulus onset and offset; (8) distance to the nest at the end of the trial; (9) total distance traveled at speeds greater than 90% of the peak speed; (10) duration of movement at speeds greater than 90% of the peak speed; and (11) duration of time with both key points moving slower than 1 cm / s. Eight features were related to state transitions: (12) latency to flee, defined as the time from stimulus onset to the first escape state, where an escape state was identified when the animal moved more than 10 cm at a speed exceeding 10 cm / s; (13-15) average speed, distance traveled, and acceleration during the fastest escape episode; (16) latency to the fastest escape episodes; (17) latency to the first stationary state, where a stationary state was defined when both key points moved less than 1 cm / s for more than 0.3 s; (18,19) total and longest stationary duration.

Revised version: Third, we defined 19 behavioral features related to locomotion speed, distance, and state transitions. Eleven features were related to speed and distance: (1) peak speed toward the reward port before stimulus onset; (2–4) maximum, mean, and coefficient of variation of speed after stimulus onset; (5) latency to peak speed; (6) hiding latency, defined as the time between stimulus onset and arrival at the safe zone (set to 20 s if the mouse did not reach the safe zone by the end of the trial); (7) escape distance, defined as the displacement from stimulus onset to offset; (8) distance to the nest at the end of the trial; (9) total distance traveled at speeds greater than 90% of the peak speed; (10) duration of movement at speeds greater than 90% of the peak speed; and (11) duration of time with both key points moving slower than 1 cm/s. Eight features were related to state transitions: (12) latency to flee, defined as the time from stimulus onset to the first escape episode, where an escape state was defined when the animal moved more than 10 cm at a speed exceeding 10 cm/s; (13–15) average speed, distance traveled, and acceleration during the fastest escape episode; (16) latency to the fastest escape episode; (17) latency to the first stationary episode, where a stationary state was defined when both key points moved at less than 1 cm/s for more than 0.3 s; (18) total stationary duration and (19) longest stationary duration.

Earlier version: Finally, these features were fed into a random forest model with a maximum depth of 5 to classify animal behaviors. For trials in which neither escape nor stationary states were identified, latencies (feature 12,16,17) were set to 20 s, distance (14) to 0 cm, speed (13) to 0 cm / s, acceleration (15) to 0 cm / s2, and duration (18,19) to 0 s. The model was trained on 238 trials and tested on 102 trials, achieving an accuracy of 0.95 on the test set. Using this classifier, behavioral decisions in 3862 trials across 140 mice were categorized into four types: direct escape, escape after assessment, freezing, and no response (Figure S1). A decision score was calculated for each condition as S = 3pe + 2pfe + pf , where pe, pae, pf denote the proportion of direct escape, escape after assessment, and freezing, respectively. To quantify how quickly animals recovered from a fear state, we measured recovery time. For escape trials, recovery time was defined as the hiding duration, measured from entry into the safe zone to re-entry into the arena. For freezing trials, recovery time was defined as the interval between the onset of the first freezing episode and the termination of the last freezing episode.

Revised version: Finally, these features were fed into a random forest model with a maximum depth of 5 to classify behaviors. For trials in which neither escape nor stationary states were identified, latencies (features 12, 16, 17) were set to 20 s, distance (14) to 0 cm, speed (13) to 0 cm/s, acceleration (15) to 0 cm/s2, and duration (18, 19) to 0 s. The model was trained on 238 trials and tested on 102 trials, achieving an accuracy of 0.95 on the test set. Using this classifier, behavioral decisions in 3862 trials across 140 mice were categorized into four types: direct escape, escape after assessment, freezing, and no response (Figure S1). A decision score was calculated for each condition as S = 3pe + 2pae + pf, where pe, pae, pf denote the proportions of direct escape, escape after assessment, and freezing, respectively. To quantify how quickly animals recovered from a fear state, we measured recovery time. For escape trials, recovery time was defined as the hiding duration, measured from entry into the safe zone to re-entry into the arena. For freezing trials, recovery time was defined as the interval between the onset of the first freezing episode and the termination of the last freezing episode.

Earlier version: To quantify the influence of reward, the reward zone was defined as a region within 10 cm of the reward port, and the duration in the reward zone was defined as the time spent within this zone during the 20 seconds following stimulus onset. Vigilance was assessed using latency to flee, foraging interval (time between two consecutive entries into the trigger zone), and foraging speed (average speed heading to the trigger zone before the looming exposure).

Revised version: To quantify the influence of reward, the reward zone was defined as a region within 10 cm of the reward port, and the duration in the reward zone was defined as the time spent within this zone during the 20 s following stimulus onset. Vigilance was assessed using latency to flee, foraging interval (time between two consecutive entries into the trigger zone), and foraging speed (average speed while approaching the trigger zone before stimulus onset).

Earlier version: To segment the trials into two phases, we performed principal component analysis on a feature matrix containing decision score, escape distance, time spent in the reward zone, peak escape speed, and latency to flee across the first 10 trials. We extracted the first principal component (Figure S3), which captured the majority of the variance, and used its learning curve to detect transitions in behavior. Change points of the learning curve were detected by the ruptures approach (Gallistel et al., 2004; Truong et al., 2020).

Revised version: To segment the trials into two phases, we performed principal component analysis on a feature matrix containing decision score, escape distance, duration in the reward zone, peak escape speed, and latency to flee across the first 10 trials. We extracted the first principal component (Figures S3C–D), which captured the majority of the variance, and used its learning curve to detect transitions in behavior. Change points in the learning curve were detected using the ruptures method (Gallistel et al., 2004; Truong et al., 2020).

Behavioral modeling

Earlier version: Parameters in the drift-diffusion leaky integrator model were estimated in two steps by minimizing a loss function using a grid search. In the first step, we fit our model to behavioral data in the no-reward condition, setting r = 0 and β = 1 for the low-threat condition, and estimated α, δ, xthr, and β for the high-threat condition. In the second step, we estimated β and r for water and sucrose conditions under both threat conditions. The loss function combines fitting errors in the distribution of behavioral decisions (ℒdec), the median latency to flee (ℒmed), and the variability of latency (ℒsd): The decision-distribution mismatch is quantified using a mean cosine distance: where dc and d^c denote the observed and model-predicted decision distributions for contrast level c, and C = 2 is the number of threat conditions.

Revised version: Parameters in the drift-diffusion leaky integrator model were estimated in two steps by minimizing a loss function using a grid search. In the first step, we fit our model to behavioral data from the no-reward condition, setting r = 0 and β = 1 for the low-threat condition, and estimated α, δ, xthr, and β for the high-threat condition. In the second step, we estimated β and r for the water and sucrose conditions under both threat conditions. The loss function combines fitting errors in the distribution of behavioral decisions (ℒdec), the median latency to flee (ℒmed), and the variability of latency (ℒsd): (3) L=100Ldec+2Lmed+Lsd The decision-distribution mismatch is quantified using a mean cosine distance: (4) Ldec=1C∑c=1C(1−dc⋅d^c‖dc‖‖d^c‖) where dc and ̂dc denote the observed and model-predicted decision distributions for threat level c, and C = 2 is the number of threat conditions.

Earlier version: The median-latency loss is defined as and the latency-variability loss is For the final fitted model (Equation 1), α = 1.78 and δ = 5.6. Under low-threat conditions, β = 1 for all conditions; under high-threat conditions, β = 1.44, 1.8, 2.55 for no reward, water, and sucrose, respectively. The reward values were r = 0, 0.08, 0.25 for no reward, water, and sucrose, respectively, under both threat conditions. In Equation 2, the decision threshold was xthr = 0.71.

Revised version: The median-latency loss is defined as (5) Lmed=1C∑c=1C|L^cmed−Lcmed|L^cmed+Lcmed, and the latency-variability loss is (6) Lsd=1C∑c=1C|L^csd−Lcsd|L^csd+Lcsd. For the final fitted model (Equation 1), α = 1.78 and δ = 5.6. Under low-threat conditions, β = 1 for all reward conditions; under high-threat conditions, β = 1.44, 1.8, 2.55 for no reward, water, and sucrose, respectively. The reward values were r = 0, 0.08, 0.25 for no reward, water, and sucrose, respectively, under both threat conditions. In Equation 2, the decision threshold was xthr = 0.71.

Earlier version: To further understand how threat and reward shape decisions, we analyze a deterministic simplification of the model: where p > 0 denotes the perceived reward value and n · t < 0 denotes the perceived threat strength, which grows linearly over time.

Revised version: To further understand how threat and reward shape decisions, we analyzed a deterministic simplifica-tion of the model: (7) dxdt=−α⋅x(t)+n⋅t−p where p > 0 denotes the perceived reward value and n · t < 0 denotes the perceived threat strength, which grows linearly over time.

Earlier version: With the initial condition x(0) = 0, the solution is This reveals a competition between two components. First, a linear term driven by threat accumulation promotes escape over time. Second, an exponentially decaying component reflects the combined influence of reward and threat.

Revised version: With the initial condition x(0) = 0, the solution is (8) x(t)=nαt+(pα+nα2)(e−αt−1) This solution reveals a competition between two components. First, a linear term driven by threat accumulation promotes escape over time. Second, an exponentially decaying component reflects the combined influence of reward and threat.

Quantification and statistical analysis

Earlier version: No statistical method was used to predetermine sample size. The Shapiro-Wilk test was applied to assess the normality of data distributions. For comparison between two groups, a two-sample t-test or paired t-test was used for normally distributed data; otherwise, the Mann–Whitney U-test or one-sample Wilcoxon signed-rank test was applied. For multi-group comparisons, one-way or two-way analysis of variance (ANOVA) was used when data were normally distributed, followed by Tukey’s post hoc test. For non-parametric multi-group comparisons, the Kruskal–Wallis test (for one factor) or the Scheirer–Ray–Hare test (for two factors) was applied, followed by Dunn’s post hoc test with Holm correction. The chi-square test was conducted to analyze categorical variables, while paired categorical comparisons were assessed using the Stuart–Maxwell test. Detailed statistical information for each experiment is provided in the Results section and figure legends.

Revised version: No statistical method was used to predetermine sample size. The Shapiro–Wilk test was applied to assess the normality of data distributions. For comparison between two groups, a two-sample t-test or paired t-test was used for normally distributed data; otherwise, the Mann–Whitney–Wilcoxon or one-sample Wilcoxon signed-rank test was applied. For multi-group comparisons, one-way or two-way analysis of variance (ANOVA) was used when data were normally distributed, followed by Tukey’s post hoc test. For non-parametric multi-group comparisons, the Kruskal–Wallis test (for one factor) or the Scheirer–Ray–Hare test (for two factors) was applied, followed by Dunn’s post hoc test with Holm correction. The chi-square test was conducted to analyze categorical variables, while paired categorical comparisons were assessed using the Stuart–Maxwell test. Detailed statistical information for each experiment is provided in the Results section and figure legends.